```python
from autograder_term2week9 import *
_globals = globals()
from copy import copy
```

# <center> Introduction to Mathematical Computing </center>
### <center>Phil Ramsden, Boumediene Hamzi, Sam Brzezicki, Matt Woolway</center>

# <center> Worksheet 11: An Introduction to Algorithms and Complexity</center>

**NOTE:** When answering a question, remember to remove the `raise NotImplementedError()` line.

This worksheet is all about different ways of sorting data. Often, we will be using the reference data list

```python
masterData = [3, 34, 12, 22, 27, 17, 31, 29, 40, 24, 21, 19, 7, 18, 26, 5]
```

of length 16.

We'll be measuring <b>complexity</b> by counting, for each algorithm, the number of <b>comparisons between data items</b>. This is an imperfect measure, because sorting algorithms also include other things, such as moving data around! But the details of that can be language-dependent, whereas the number of comparisons is what it is, and is independent of language.

## Question 1

This question is about <b>bubble sort</b>, one of the simplest sorting algorithms of all. For a data list of length 16, bubble sort proceeds as follows:
<ul>
    <li>Compare items 0 and 1, and if necessary swap them. Then compare items 1 and 2, again swapping if necessary; then 2 and 3, and so on up to 14 and 15. You can now guarantee the maximum element is in the 15 position.</li>
    <li>Compare items 0 and 1, and if necessary swap them. Then compare items 1 and 2, again swapping if necessary; then 2 and 3, and so on up to 13 and 14. You can now guarantee the maximum element is in the 14 position.</li>
    <li>...</li>
    <li>Compare items 0 and 1, and if necessary swap them. Then compare items 1 and 2, again swapping if necessary.</li>
    <li>Compare items 0 and 1, and if necessary swap them.</li>
</ul>

Animations are available to view on Blackboard.

### (i) 

Here is a listing of a simple bubble sort function.

```python
def bubble_sort_simple(data):
    """
    Sorts data using the bubble sort algorithm, with no tracking
    of orderedness
    """
    
    # number of data items
    ndata = len(data)
    
    # end index goes down from ndata to 2
    for end_index in range(ndata,1,-1):
        # comparison index goes up from 0 to end index - 2
        for comparison_index in range(0, end_index-1):
            # if element is greater than its neighbour to the right, swap them
            if data[comparison_index] > data[comparison_index+1]:
                data[comparison_index], data[comparison_index+1] = \
                data[comparison_index+1], data[comparison_index]
```

The function sorts data <em>in place</em>: that is, it returns nothing, but has the side-effect of sorting the list given to it as its argument.

(a) Use `copy`, imported above from the `copy` module, to create a copy of `masterData` called `data`.

<!-- grade_id: cell-7f9747d7b626e328, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
data = copy(masterData)
### END SOLUTION

print(data)
```

<!-- grade_id: cell-5897e6f73013494c, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1ia(_globals)
```

(b) Test `bubble_sort_simple` by sorting `data`.

<!-- grade_id: cell-d9492ddc080ee000, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
bubble_sort_simple(data)
### END SOLUTION

print(data)
```

<!-- grade_id: cell-33a351855e2a8177, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1ib(_globals)
```

Here is a listing of `bubble_sort_simple` to which a global variable called `comparison_count` has been added, which increments by 1 every time a comparison is made.

```python
def bubble_sort_simple(data):
    """
    Sorts data using the bubble sort algorithm, with no tracking
    of orderedness
    """
    
    # comparison count
    global comparison_count
    
    # number of data items
    ndata = len(data)
    
    # end index goes down from ndata to 2
    for end_index in range(ndata,1,-1):
        # comparison index goes up from 0 to end index - 2
        for comparison_index in range(0, end_index-1):
            # if element is greater than its neighbour to the right, swap them
            comparison_count += 1
            if data[comparison_index] > data[comparison_index+1]:
                data[comparison_index], data[comparison_index+1] = \
                data[comparison_index+1], data[comparison_index]
```

(c) Define `data` as a copy of `masterData` again. Then set `comparison_count` to 0, and run `bubble_sort_simple`. Finally, check the new value of `comparison_count`.

<!-- grade_id: cell-3f0ba9a511d2657a, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
data = copy(masterData)
comparison_count = 0
bubble_sort_simple(data)
### END SOLUTION

print(data)
print(comparison_count)
```

<!-- grade_id: cell-9745abfd9ac67e65, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1ic(_globals)
```

(d) Experiment with sorting other data sets of length 16. If the value of `comparison_count` starts at zero each time, what is its final value?

<ol>
    <li>With a data set of length 16, the final value of <code>comparison_count</code> is always exactly 120.</li>
    <li>With a data set of length 16, the final value of <code>comparison_count</code> is always at least 120, and sometimes more than that.</li>
    <li>With a data set of length 16, the final value of <code>comparison_count</code> is always at most 120, and sometimes less than that.</li>
    <li>With a data set of length 16, the final value of <code>comparison_count</code> is sometimes less than 120, sometimes more than 120, and sometimes exactly 120.</li>
</ol>

<!-- grade_id: cell-f8040985612b2374, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
q1id_answer = 1
### END SOLUTION

q1id_answer = # insert one of 1, 2, 3 or 4
```

<!-- grade_id: cell-69343024e1dfbe39, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(q1id_answer == question1id())
print('test case passed!')
```

(e) Experiment with data sets of different lengths.

Then, in the form of lambda-expressions, state the maximum and minimum number of comparisons carried out by `bubble_sort_simple` when the data set is of length `n`. So, for example, if you think the maximum number of comparisons is $n^2+4$ and the minimum is $n$, you'd type
```python
bubble_sort_simple_max = lambda n: n**2+4
bubble_sort_simple_min = lambda n: n
```
Try to make sure the output from each of your lambda-expressions is always an int.

<!-- grade_id: cell-4a32972f7163af8d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
bubble_sort_simple_max = lambda n: n*(n-1)//2
bubble_sort_simple_min = lambda n: n*(n-1)//2
### END SOLUTION

# quick test
print(bubble_sort_simple_max(16))
print(bubble_sort_simple_min(16))
```

<!-- grade_id: cell-0f48fb8a641a5cec, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1ie(_globals)
```

(f) Which of the following is true?

<ol>
    <li><code>bubble_sort_simple</code> is $O(n^2)$ and $\Omega(n^2)$ and is therefore $\Theta(n^2)$;</li>
    <li><code>bubble_sort_simple</code> is $O(n^2)$ but not $\Omega(n^2)$;</li>
    <li><code>bubble_sort_simple</code> is $\Omega(n^2)$ but not $O(n^2)$;</li>
    <li><code>bubble_sort_simple</code> is neither $O(n^2)$ nor $\Omega(n^2)$.</li>
</ol>

<!-- grade_id: cell-08fd347e97ef866c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
q1if_answer = 1
### END SOLUTION

q1if_answer = # insert one of 1, 2, 3 or 4
```

<!-- grade_id: cell-382e7382c2787323, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(q1if_answer == question1if())
print('test case passed!')
```

### (ii)

(a) Write a function called `bubble_sort_enhanced` that also sorts data in place using bubble sort, but with the following additional feature. 

On each pass through the data (that is, on each turn of the outer `for` loop), a record is kept of whether any swaps were necessary; if not, the data is regarded as sorted, and the process stops. (So if `data` were already sorted, only one pass would be necessary.)

Make sure to give your function a global variable called `comparison_count` that goes up by 1 every time a comparison is made.

<!-- grade_id: cell-2296e52df6b80511, points: 0, grade: false, solution: true -->
```python
def bubble_sort_enhanced(data):
    ### BEGIN SOLUTION
    """
    Sorts data using the bubble sort algorithm, with tracking
    of orderedness
    """
    
    # comparison count
    global comparison_count
    
    # number of data items
    ndata = len(data)
    
    # orderedness tracking
    already_ordered = False
    
    # end index goes down from ndata to 2
    for end_index in range(ndata,1,-1):
        if already_ordered:
            break
        already_ordered = True
        # comparison index goes up from 0 to end index - 2
        for comparison_index in range(0, end_index-1):
            # if element is greater than its neighbour to the right, swap them
            comparison_count += 1
            if data[comparison_index] > data[comparison_index+1]:
                already_ordered = False
                data[comparison_index], data[comparison_index+1] = \
                data[comparison_index+1], data[comparison_index]
    ### END SOLUTION
```

<!-- grade_id: cell-1d764dd6ddc55bd5, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1iia(_globals)
```

(b) Test your function on a copy of `masterData`.

<!-- grade_id: cell-66a9f56a544610a1, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
data = copy(masterData)
comparison_count = 0
bubble_sort_enhanced(data)
### END SOLUTION

print(data)
print(comparison_count)
```

<!-- grade_id: cell-78f4dc871de0fc7a, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1iib(_globals)
```

(c) Experiment with your new function, using data sets of length 16. If the value of `comparison_count` starts at zero each time, what is its final value?

<ol>
    <li>With a data set of length 16, the final value of <code>comparison_count</code> is always exactly 120.</li>
    <li>With a data set of length 16, the final value of <code>comparison_count</code> is always at least 120, and sometimes more than that.</li>
    <li>With a data set of length 16, the final value of <code>comparison_count</code> is always at most 120, and sometimes less than that.</li>
    <li>With a data set of length 16, the final value of <code>comparison_count</code> is sometimes less than 120, sometimes more than 120, and sometimes exactly 120.</li>
</ol>

<!-- grade_id: cell-e4e02775dd44384c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
q1iic_answer = 3
### END SOLUTION

q1iicanswer = # insert one of 1, 2, 3 or 4
```

<!-- grade_id: cell-3015963ad8f74893, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(q1iic_answer == question1iic())
print('test case passed!')
```

(d) Experiment with data sets of different lengths. Then, in the form of lambda-expressions, state the maximum and minimum number of comparisons carried out by `bubble_sort_enhanced` when the data set is of length `n`. So, for example, if you think the maximum number of comparisons is $n^2+4$ and the minimum is $n$, you'd type
```python
bubble_sort_enhanced_max = lambda n: n**2+4
bubble_sort_enhanced_min = lambda n: n
```

<!-- grade_id: cell-8393bed1a818446e, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
bubble_sort_enhanced_max = lambda n: n*(n-1)//2
bubble_sort_enhanced_min = lambda n: n - 1
### END SOLUTION

# quick test
print(bubble_sort_enhanced_max(16))
print(bubble_sort_enhanced_min(16))
```

<!-- grade_id: cell-e34997ca4e968f1c, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1iid(_globals)
```

(e) Which of the following is true?

<ol>
    <li><code>bubble_sort_enhanced</code> is $O(n^2)$ and $\Omega(n^2)$ and is therefore $\Theta(n^2)$;</li>
    <li><code>bubble_sort_enhanced</code> is $O(n^2)$ but not $\Omega(n^2)$;</li>
    <li><code>bubble_sort_enhanced</code> is $\Omega(n^2)$ but not $O(n^2)$;</li>
    <li><code>bubble_sort_enhanced</code> is neither $O(n^2)$ nor $\Omega(n^2)$.</li>
</ol>

<!-- grade_id: cell-e3da7eed97a107a2, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
q1iie_answer = 2
### END SOLUTION

q1iieanswer = # insert one of 1, 2, 3 or 4
```

<!-- grade_id: cell-8345e676a61b091f, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(q1iie_answer == question1iie())
print('test case passed!')
```

## Question 2

This question is about <b>insertion sort</b>, which proceeds as follows.

At every stage, think of $a_0$ as the item in position 0, $a_1$ the item in position 1, and so on; the values of these $a$'s may change.

<ul>
    <li>Compare $a_1$ with $a_0$. If $a_1 \ge a_0$, do nothing; the pass is finished. 
        
Otherwise, insert $a_1$ in position $0$ and shift $a_0$ along to position 1. 

The list is now sorted from position 0 to position 1.</li>
    <li>Compare $a_2$ with (the possibly new) $a_1$. If $a_2 \ge a_1$, do nothing. Otherwise, compare $a_2$ with $a_0$. 
    
If $a_2 \ge a_0$, insert $a_2$ in position 1 and shift $a_1$ along to position 2; otherwise, insert $a_2$ in position 0 and shift $a_0$ and $a_1$ along to positions 1 and 2 respectively.  

The list is now sorted from position 0 to position 2.</li>
    <li>Compare $a_3$ with (the possibly new) $a_2$. If $a_3 \ge a_2$, do nothing. Otherwise, compare $a_3$ with $a_1$. 
    
If $a_3 \ge a_1$, insert $a_3$ in position 2 and shift $a_2$ along to position 3; otherwise, compare $a_3$ with $a_0$. 

If $a_3 \ge a_0$, insert $a_3$ in position 1 and shift $a_1$ and $a_2$ along to positions 2 and 3; otherwise, insert $a_3$ in position 0 and shift $a_0$, $a_1$ and $a_2$ along to positions 1, 2 and 3.  

The list is now sorted from position 0 to position 3.</li>
    <li>And so on: on each pass, compare $a_n$ with $a_{n-1}$, then with $a_{n-2}$, etc, until the correct position for $a_n$ is found; then insert it there, moving as many items to the right as necessary to make room for it.

The list will then be sorted from position 0 to position $n$.

Keep going until the final element.</li>
</ul>

You can view animations on Blackboard.

(a) Write a function called `insertion_pass` which should take two arguments: a list, `data`, and a positive integer representing the length of the already sorted sublist, `already_sorted`. It should then make one pass of insertion sort, inserting the next element in the correct position and increasing the length of the already sorted sublist by one.

Thus for example
```python
data = [3, 5, 9, 4, 7, 2, 1, 6]
insertion_pass(data, 3)
print(data)
```
should give `[3, 4, 5, 9, 7, 2, 1, 6]`.

Take care that your function only performs such comparisons as are necessary, stopping as soon as it can. Do not at this stage keep a comparison count, though.

<!-- grade_id: cell-03416a7da53b1633, points: 0, grade: false, solution: true -->
```python
def insertion_pass(data, already_sorted):
    ### BEGIN SOLUTION
    """
    Performs one pass of insertion sort, lengthening the already sorted
    sublist by one.
    """
    
    # next item to be inserted
    next_item = data[already_sorted]
    
    # initial value of insertion index
    insertion_index = already_sorted-1
    
    # while loop
    while next_item < data[insertion_index] and insertion_index > -1:
        insertion_index -= 1
    
    # insert to the right of insertion_index
    data[insertion_index+2:already_sorted+1] = \
    data[insertion_index+1:already_sorted]
    data[insertion_index+1] = next_item
    ### END SOLUTION

# quick test
data = [3, 5, 9, 4, 7, 2, 1, 6]
insertion_pass(data, 3)
print(data)
```

<!-- grade_id: cell-d29ef87c4e42e6f8, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2a(_globals)
```

(b) Hence, write a function called `insertion_sort`, which sorts `data` in place using insertion sort.

<!-- grade_id: cell-39fc3085c2b00537, points: 0, grade: false, solution: true -->
```python
def insertion_sort(data):
    ### BEGIN SOLUTION
    """
    Sorts data using insertion sort
    """
    
    for already_sorted in range(1,len(data)):
        insertion_pass(data, already_sorted)
    ### END SOLUTION
```

<!-- grade_id: cell-70e6a78021d7268b, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2b(_globals)
```

(c) Test your new function on `data`, which should be a copy of `masterData`.

<!-- grade_id: cell-db412796a141c807, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
data = copy(masterData)
insertion_sort(data)
### END SOLUTION

print(data)
```

<!-- grade_id: cell-44dd959ce673b63d, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2c(_globals)
```

(d) Try to think this through from first principles! With a data list of length 16, what is the maximum number of comparisons that are necessary for insertion sort, if on every pass it's necessary to compare $a_n$ with $a_{n-1}$, $a_{n-2}$, all the way down to $a_0$?

<!-- grade_id: cell-a0106c641a3c6e1f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
q2d_answer = 120
### END SOLUTION

q2d_answer = # insert your answer here as an int
```

<!-- grade_id: cell-841d490288aa8680, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(q2d_answer == question2d())
print('test case passed!')
```

(e) Give an example of a "worst case" data list of length 16 for which this many comparisons are necessary.

<!-- grade_id: cell-cdbd5264310c1038, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
insertion_sort_length16_worstcase = list(range(16,0,-1))
### END SOLUTION

#insertion_sort_length16_worstcase = None# insert your example
```

<!-- grade_id: cell-31776121b072bf85, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2e(_globals)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-2-82ea22c13804> in <module>
      3 # Run this cell for grading
      4 _globals = globals()
----> 5 question2e(_globals)

NameError: name 'question2e' is not defined
```

(f) Try to think this through from first principles! With a data list of length 16, what is the minimum number of comparisons that are necessary for insertion sort, if on every pass it's only necessary to compare $a_n$ with $a_{n-1}$?

<!-- grade_id: cell-691ac8cf4516f2c7, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
q2f_answer = 15
### END SOLUTION

q2f_answer = # insert your answer here as an int
```

<!-- grade_id: cell-b566cccf4417a613, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(q2f_answer == question2f())
print('test case passed!')
```

(g) Give an example of a "best case" data list of length 16 for which this many comparisons are necessary.

<!-- grade_id: cell-db849fdbc2093719, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
insertion_sort_length16_bestcase = list(range(16))
### END SOLUTION

insertion_sort_length16_bestcase = # insert your example here
```

<!-- grade_id: cell-1ea7a4fe94aef220, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2g(_globals)
```

(h) Introduce a global variable called `comparison_count` into your `insertion_pass` function, which should go up by 1 for each comparison it makes. (This might need some thought; depending on how your program works, it could be quite tricky.) 

Once you've done this, you will have no need to make any changes to `insertion_sort`; the count of comparisons will happen automatically.

Use this feature to test your answers to (c), (d), (e) and (f).

Test this feature too on a copy of `masterData`.

<!-- grade_id: cell-d60f031ae23949de, points: 0, grade: false, solution: true -->
```python
# Cell for defining your function

def insertion_pass(data, already_sorted):
    ### BEGIN SOLUTION
    """
    Performs one pass of insertion sort, lengthening the already sorted
    sublist by one.
    """
    
    global comparison_count
    
    # next item to be inserted
    next_item = data[already_sorted]
    
    # initial value of insertion index
    insertion_index = already_sorted-1
    
    # while loop
    comparison_count += 1
    while next_item < data[insertion_index] and insertion_index > -1:
        comparison_count += 1
        insertion_index -= 1
    
    if insertion_index == -1:
        comparison_count -= 1
    
    # insert to the right of insertion_index
    data[insertion_index+2:already_sorted+1] = \
    data[insertion_index+1:already_sorted]
    data[insertion_index+1] = next_item
    ### END SOLUTION
```

<!-- grade_id: cell-0c5de2505f5a39ea, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2hfunc(_globals)
```

<!-- grade_id: cell-07bb149de176d2bf, points: 0, grade: false, solution: true -->
```python
# cell for testing on a copy of masterData

### BEGIN SOLUTION
data = copy(masterData)
comparison_count = 0
insertion_sort(data)
### END SOLUTION

print(data)
print(comparison_count)
```

<!-- grade_id: cell-e0a8eaee2d129964, points: 2, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2htest(_globals)
```

(j) Experiment. Which of the following seems to be true?

<ol>
    <li>Insertion sort outperforms bubble sort in the worst case and the best case, and seems to outperform it in the typical case.</li>
    <li>Insertion sort outperforms bubble sort in the worst case, and seems to outperform it in the typical case, but performs exactly as well in the best case.</li>
    <li>Insertion sort outperforms bubble sort in the best case, and seems to outperform it in the typical case, but performs exactly as well in the worst case.</li>
    <li>Insertion sort seems to outperform bubble sort in the typical case, but performs exactly as well in the worst case and best case.</li>
    <li>Insertion sort performs exactly as well as bubble sort in the worst case and best case, and roughly as well in the typical case.</li>
</ol>

<!-- grade_id: cell-e6188a6359fa00ca, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
q2j_answer = 4
### END SOLUTION

q2j_answer = # insert one of 1, 2, 3, 4 or 5
```

<!-- grade_id: cell-99c353404a5c658e, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(q2j_answer == question2j())
print('test case passed!')
```

## Question 3

This question is about <b>merge sort</b>, which is a more sophisticated sorting algorithm than selection sort, bubble sort or insertion sort.

### (i)

At its core is the idea of merging the contents of two <b>buffers</b>, each of which is assumed to be itself already sorted. Before we get on to how the sort itself works, let's consider this merging operation.

We start with two lists of length 3 and 4 respectively, which we'll call `left_buffer` and `right_buffer`, and a destination list called `data`, of length 7. Now, `data` can have any old junk in it as far as the merge operation is concerned, so let's fill it with zeros. So we start off like this:

`left_buffer`: 3, 7, 8

`right_buffer`: 2, 5, 9, 10

`data`: 0, 0, 0, 0, 0, 0, 0

Notice that `left_buffer` and `right_buffer` are already sorted.

Our first move is to compare the position 0 elements of the two buffers: 3 and 2. The smaller is 2, so we choose that one, and place it in position 0 of data:

`left_buffer`: 3, 7, 8

`right_buffer`: <font color='gray'>2</font>, 5, 9, 10

`data`: 2, 0, 0, 0, 0, 0, 0

I've "grayed out" 2, as it's been used.

Now we compare the position 0 element of `left_buffer` with the position 1 element of the partially used-up `right_buffer`: 3 and 5. The smaller is 3, so we use that:

`left_buffer`: <font color='gray'>3</font>, 7, 8

`right_buffer`: <font color='gray'>2</font>, 5, 9, 10

`data`: 2, 3, 0, 0, 0, 0, 0

Now we compare the position 1 elements in each buffer: 5 and 7. This gives

`left_buffer`: <font color='gray'>3</font>, 7, 8

`right_buffer`: <font color='gray'>2</font>, <font color='gray'>5</font>, 9, 10

`data`: 2, 3, 5, 0, 0, 0, 0

And so on. We eventually get to this situation:

`left_buffer`: <font color='gray'>3</font>, <font color='gray'>7</font>, <font color='gray'>8</font>

`right_buffer`: <font color='gray'>2</font>, <font color='gray'>5</font>, 9, 10

`data`: 2, 3, 5, 7, 8, 0, 0

At this point, with one of the buffers used up, we're done, and we can fill in the remaining data:

`left_buffer`: <font color='gray'>3</font>, <font color='gray'>7</font>, <font color='gray'>8</font>

`right_buffer`: <font color='gray'>2</font>, <font color='gray'>5</font>, <font color='gray'>9</font>, <font color='gray'>10</font>

`data`: 2, 3, 5, 7, 8, 9, 10


(a) Write a function called `merge`, which takes as its arguments three lists, `left_buffer`, `right_buffer` and `data`, and merges `left_buffer` and `right_buffer` into `data` in place (that is, as a side-effect; it should return no output).

Try to make sure you carry out only such comparisons as are necessary, but don't at this stage include a comparison count.

Test it: the following code
```python
left_buffer = [3, 7, 8]
right_buffer = [2, 5, 9, 10]
data = [0, 0, 0, 0, 0, 0, 0]
merge(left_buffer, right_buffer, data)
print(data)
```

should give `[2, 3, 5, 7, 8, 9, 10]`

<!-- grade_id: cell-5e22b2fe4863815c, points: 0, grade: false, solution: true -->
```python
def merge(left_buffer, right_buffer, data):
    ### BEGIN SOLUTION
    """
    Merges left_buffer and right_buffer into data in place
    """
    
    # initialize indexes
    left_index = right_index = main_index = 0
    
    # loop for as long as neither buffer used up
    while left_index < len(left_buffer) and right_index < len(right_buffer):
        # left buffer has smaller element
        if left_buffer[left_index] < right_buffer[right_index]:
            data[main_index] = left_buffer[left_index]
            left_index += 1
        # otherwise
        else:
            data[main_index] = right_buffer[right_index]
            right_index += 1
        main_index += 1
    
    # left buffer not used up
    if left_index < len(left_buffer):
        data[main_index:] = left_buffer[left_index:]
    # right buffer not used up
    else:
        data[main_index:] = right_buffer[right_index:]
    ### END SOLUTION
    
# quick test
left_buffer = [3, 7, 8]
right_buffer = [2, 5, 9, 10]
data = [0, 0, 0, 0, 0, 0, 0]
merge(left_buffer, right_buffer, data)
print(data)
```

<!-- grade_id: cell-02a6bb6d95db6df1, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3ia(_globals)
```

(b) What is the least number of comparisons necessary in order to merge two sorted buffers of length 8 into a data list of length 16?

<!-- grade_id: cell-8ed0c9720cee139e, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
q3ib_answer = 8
### END SOLUTION

q3ib_answer = # insert your answer here as an int
```

<!-- grade_id: cell-433aefcb80f0a160, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(q3ib_answer == question3ib())
print('test case passed!')
```

(c) What is the greatest number of comparisons necessary in order to merge two sorted buffers of length 8 into a data list of length 16?

<!-- grade_id: cell-724c6fa937558218, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
q3ic_answer = 15
### END SOLUTION

q3ic_answer = # insert your answer here as an int
```

<!-- grade_id: cell-c4cd301632439d83, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(q3ic_answer == question3ic())
print('test case passed!')
```

### (ii) 

One way to implement merge sort itself is <b>recursively</b>. The algorithm goes like this.

To merge sort `data`:

If `data` consists of only one element, do nothing. Otherwise,

<ul>
    <li>Let <code>half_length</code> be half the length of <code>data</code>.</li>
    <li>Let <code>left_half</code> be <code>data[0:half_length]</code> and <code>right_half</code> the rest.</li>
    <li>Merge sort <code>left_half</code>.</li>
    <li>Merge sort <code>right_half</code>.</li>    
    <li>Merge <code>left_half</code> and <code>right_half</code> into <code>data</code>.</li>
</ul>

(a) Write a function called `merge_sort` based on the above description, and test it on a variety of data lists.

<!-- grade_id: cell-0a0b52801a29d664, points: 0, grade: false, solution: true -->
```python
def merge_sort(data):
    ### BEGIN SOLUTION
    """
    Sorts data using merge sort
    """
    
    # number of data items
    ndata = len(data)
    
    # if data is of length 1, do nothing; otherwise...
    if ndata > 1:
        # half the length
        half_length = ndata//2
        
        # two buffers
        left_half = data[0:half_length]
        right_half = data[half_length:ndata]
        
        # recursively sort the buffers
        merge_sort(left_half)
        merge_sort(right_half)
        
        # merge
        merge(left_half,right_half,data)
    ### END SOLUTION
```

<!-- grade_id: cell-048ab5d080208523, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3iia(_globals)
```

To merge sort a list of length 16, you must first merge sort two lists of length 8, then merge them; each list of length 8 has been created from two merge-sorted lists of length 4, and so on.

(b) In the best case, how many comparisons are necessary in order to merge sort a list of length 16?

<!-- grade_id: cell-1f6f920745570a96, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
q3iib_answer = 32
### END SOLUTION

q3iib_answer = # insert your answer here as an int
```

<!-- grade_id: cell-af5fbf8fd5ab73e9, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(q3iib_answer == question3iib())
print('test case passed!')
```

(c) In the worst case, how many comparisons are necessary in order to merge sort a list of length 16?

<!-- grade_id: cell-fce1317ba6e870b5, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
q3iic_answer = 49
### END SOLUTION

q3iic_answer = # insert your answer here as an int
```

<!-- grade_id: cell-5ffd72f11f725263, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(q3iic_answer == question3iic())
print('test case passed!')
```

(d) <b>This is genuinely tricky</b>. In the best case, how many comparisons are necessary in order to merge sort a list of length $2^m$? Give your answer as a lambda-expression called `merge_sort_min`.

<!-- grade_id: cell-6f326d0825f9f749, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
merge_sort_min = lambda m: m*2**(m-1)
### END SOLUTION

# quick test
print(merge_sort_min(4))
```

<!-- grade_id: cell-69ef9e70d14bf4bc, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3iid(_globals)
```

(e) <b>This is also genuinely tricky</b>. In the worst case, how many comparisons are necessary in order to merge sort a list of length $2^m$? Give your answer as a lambda-expression called `merge_sort_max`.

<!-- grade_id: cell-bb524909efbae5a1, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
merge_sort_max = lambda m: (m-1)*2**m+1
### END SOLUTION

# quick test
print(merge_sort_max(4))
```

<!-- grade_id: cell-dfe1f9828a3d0f07, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3iie(_globals)
```

(f) Which of the following is true:

<ol>
    <li>Merge sort is $\Theta(n)$.</li>
    <li>Merge sort is $\Theta(n\,\log n)$.</li>
    <li>Merge sort is $\Theta(n^2)$.</li>
    <li>Merge sort is $\Theta(2^n)$.</li>
    <li>Merge sort is $\Theta(n\times2^n)$.</li>
    <li>None of the above.</li>
</ol>

<!-- grade_id: cell-06bca642bce2ff30, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
q3iif_answer = 2
### END SOLUTION

q3iif_answer = # insert one of 1, 2, 3, 4 or 5
```

<!-- grade_id: cell-e4c65eaa0651265b, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(q3iif_answer == question3iif())
print('test case passed!')
```

### (iii) 

(a) Introduce a `comparison_count` into your `merge` function, and use it to test your answers to (b), (c), (d), (e) and (f) above. (Autograding won't help with all that testing, sadly, but please do carry it out.)

Be sure only to count actual comparisons of data items (and not, for example, index checks to test whether a buffer has been used up yet).

<!-- grade_id: cell-5e783663f1cad858, points: 0, grade: false, solution: true -->
```python
# cell for defining your function

def merge(left_buffer, right_buffer, data):
    ### BEGIN SOLUTION
    """
    Merges left_buffer and right_buffer into data in place
    """
    
    global comparison_count
    
    # initialize indexes
    left_index = right_index = main_index = 0
    
    # loop for as long as neither buffer used up
    while left_index < len(left_buffer) and right_index < len(right_buffer):
        # left buffer has smaller element
        comparison_count += 1
        if left_buffer[left_index] < right_buffer[right_index]:
            data[main_index] = left_buffer[left_index]
            left_index += 1
        # otherwise
        else:
            data[main_index] = right_buffer[right_index]
            right_index += 1
        main_index += 1
    
    # left buffer not used up
    if left_index < len(left_buffer):
        data[main_index:] = left_buffer[left_index:]
    # right buffer not used up
    else:
        data[main_index:] = right_buffer[right_index:]
    ### END SOLUTION
```

<!-- grade_id: cell-d4d6ef12e4d93202, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3iiiafunc(_globals)
```

<!-- grade_id: cell-9012d33cf5cfb424, points: 0, grade: false, solution: true -->
```python
# for testing on a copy of masterData
### BEGIN SOLUTION
comparison_count = 0
data = copy(masterData)
merge_sort(data)
# END SOLUTION

print(data)
print(comparison_count)
```

<!-- grade_id: cell-b751a3fcb95b26c5, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3iiiatest(_globals)
```

(b) Find an example of a data list of length 16 that needs the fewest possible comparisons using merge sort.

<!-- grade_id: cell-d88a0a548f05d8ab, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
merge_sort_length16_bestcase = list(range(16))
### END SOLUTION

merge_sort_length16_bestcase = # insert your example here
```

<!-- grade_id: cell-dfed020fc8d336a2, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3iiib(_globals)
```

(c) <b>Yet another genuinely tricky one</b>: find an example of a data list of length 16 that needs the most possible comparisons using merge sort.

<!-- grade_id: cell-df8754ecb97b3316, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
merge_sort_length16_worstcase = [1,9,5,13,3,11,7,15,2,10,6,14,4,12,8,16]
### END SOLUTION

merge_sort_length16_worstcase = # insert your example here
```

<!-- grade_id: cell-2d52fecafa3bcf81, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3iiic(_globals)
```

## Additional stuff to do if you get time.

Try out your three sorting algorithms, plus selection sort, on genuinely large data lists, such as 10000 randomly-generated floats.

Investigate `quicksort`, which you can also see an animation of: how does it work, what is its best-case and worst-case performance? Can you implement it?

Investigate `shellsort`, for which there isn't an animation (I may manage one next year, but they take a while to design and build).

Find out about Python's sort algorithm, `timsort`.

```python

```