```python
from autograder_term2week3 import *
_globals = globals()
```

Output:
```
Autograder loaded successfully!
Remember to always restart and run all from the Kernel menu before submitting!
```

# <center> Introduction to Mathematical Computing </center>
### <center>Phil Ramsden, Boumediene Hamzi, Sam Brzezicki</center>

# <center> Worksheet 6: More About Functions</center>

**NOTE:** When answering a question, remember to remove the `raise NotImplementedError()` line.

## Question 1

For a planet or comet, the <b>mean anomaly</b>, $M$, is the time since the closest approach to the Sun, divided by the length of the year, multiplied by $2\pi$ (so for example if one-quarter of the year has gone by, $M=\pi/2$). The <b>eccentric anomaly</b>, $E$, is a quantity we use to locate the planet or comet on its orbit. 
    
<b>Kepler's Equation</b> links the two. If the <b>orbital eccentricity</b> is $e$, then Kepler's Equation is
$$M=E-e\,\sin E.$$

(a) For a comet, the eccentricity $e$ can be quite large. Setting $e=0.95$ and $M=\pi/2$, estimate the value of $E$ by iterating the function
$$M+e\,\sin E$$
10 times, starting from the value $E=M$, and <em>making sure to use a lambda-expression</em>.

The function `nestlist` from the notes will be useful, and is supplied. Call your list of 11 iterates `Elist1`.

<!-- grade_id: cell-f80e92e0900ca624, points: 0, grade: false, solution: true -->
```python
from math import pi, sin, cos

def nestlist(f, x0, n):
    "Returns iterates of x = f(x), x0, x1, ... xn"
    
    # initialise variable and list
    x = x0
    xlist = [x0]
    
    # for loop
    for r in range(n):
        # update variable
        x = f(x)
        # append latest value to list
        xlist.append(x)
    
    # return list
    return xlist

M, e = pi/2, 0.95
Elist1 = nestlist(lambda E: M+e*sin(E), M, 10)
print(Elist1)
```

Output:
```
[1.5707963267948966, 2.5207963267948967, 2.1233952617855856, 2.379401380369645, 2.2267790252052713, 2.3236228059783635, 2.264067672197098, 2.301497907757926, 2.2782673649466303, 2.292803843066782, 2.283752791950874]
```

<!-- grade_id: cell-53518b4701ec028d, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1a(_globals)
```

Output:
```
Test passed!
1
```

(b) Estimate the mean anomaly again, but this time by instead iterating
$$\frac{M+e\,\sin E + E}{2}$$
10 times. Use the same eccentricity, $e=0.95$, the same mean anomaly, $M=\pi/2$ and the same starting value, $E=M$. Once more, use a lambda-expression in your call to `nestlist`. This time, call the list of 11 iterates `Elist2`.

<!-- grade_id: cell-f976fa3495d05ab2, points: 0, grade: false, solution: true -->
```python
M, e = pi/2, 0.95
Elist2 = nestlist(lambda E: (M+e*sin(E)+E)/2, M, 10)
print(Elist2)
```

Output:
```
[1.5707963267948966, 2.0457963267948966, 2.2307103695659016, 2.2760246901225174, 2.285105692211478, 2.286835289752431, 2.2871613522757923, 2.28722270121186, 2.2872342398078462, 2.287236409852908, 2.2872368179644407]
```

<!-- grade_id: cell-6b2f0edf42d38cab, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1b(_globals)
```

Output:
```
Test passed!
1
```

(c) Estimate the mean anomaly a third time, but this time by instead iterating
$$\frac{M+e\,\sin E - E\,e\,\cos E}{1- e\,\cos E}$$
10 times. Use the same eccentricity, $e=0.95$, the same mean anomaly, $M=\pi/2$ and the same starting value, $E=M$. Once more, use a lambda-expression in your call to `nestlist`. This time, call the list of 11 iterates `Elist3`.

<!-- grade_id: cell-014d16016afd87ab, points: 0, grade: false, solution: true -->
```python
M, e = pi/2, 0.95
Elist3 = nestlist(lambda E: (M+e*sin(E)-E*e*cos(E))/(1-e*cos(E)), M, 10)
print(Elist3)
```

Output:
```
[1.5707963267948966, 2.520796326794897, 2.296623574573425, 2.287256163407903, 2.287236912575581, 2.28723691249383, 2.28723691249383, 2.28723691249383, 2.28723691249383, 2.28723691249383, 2.28723691249383]
```

<!-- grade_id: cell-f304b70d40d75159, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1c(_globals)
```

Output:
```
Test passed!
1
```

(d) Write and test a function called `eccentric_anomaly`, which takes as its arguments floats $M$ and $e$, and a non-negative integer $M$, and returns, as a float, an estimate for the eccentric anomaly associated with mean anomaly $M$ and eccentricity $e$. This should be obtained from Kepler's equation by iterating $n$ times, starting with $E=M$, using the most rapidly converging of the three iterations above. You may assume that $0\le e < 1$.

<!-- grade_id: cell-0edaadf7edf15359, points: 0, grade: false, solution: true -->
```python
def eccentric_anomaly(M, e, n):
    from math import pi, sin, cos
    """
    Returns an estimate for the eccentric anomaly associated
    with mean anomaly M, and eccentricity e, obtained by n
    iterations of a rapidly converging sequence
    """
    lis = nestlist(lambda E: (M+e*sin(E)-E*e*cos(E))/(1-e*cos(E)), M, n)
    return lis[-1]
```

<!-- grade_id: cell-4c4e58d7e8c3d46e, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1d(_globals)
```

Output:
```
Testing: typical case
eccentric_anomaly(pi/2, 0.95, 5) returned the correct output

Testing: edge cases
eccentric_anomaly(pi/2, 0.95, 0) returned the correct output
eccentric_anomaly(pi/2, 0, 5) returned the correct output

Question 1d: 2 out of 2
2
```

## Question 2

(a) Write a function called <code>integer_digits</code>, which takes as its arguments <code>number</code> (an int greater than or equal to 0) and <code>base</code> (an int greater than 1), and returns, as a list, the digits of <code>number</code> in base <code>base</code>. 

Test your function: the commands <code>integer_digits(347, 10)</code> and <code>integer_digits(347, 2)</code> should return, respectively, <code>[3, 4, 7]</code> and <code>[1, 0, 1, 0, 1, 1, 0, 1, 1]</code>. If <code>number</code> is equal to zero, an empty list should be returned.

Make sure your code also works if you call it using <b>keywords</b>: for example,

```python
integer_digits(number = 347, base = 10)
```

<!-- grade_id: cell-6c432b941e2929af, points: 0, grade: false, solution: true -->
```python
from math import ceil, log
def integer_digits(number, base=10):
    lis = []
    res = number
    rem = 0
    for num in range(ceil(log(number, base))):
        res, rem = divmod(res, base)
        print(res, rem)
        lis.append(rem)
    if base == 10:
        lis.reverse()
    return lis
        
        
        
print(integer_digits(294, 2))
print(integer_digits(347, 2))
```

Output:
```
147 0
73 1
36 1
18 0
9 0
4 1
2 0
1 0
0 1
[0, 1, 1, 0, 0, 1, 0, 0, 1]
173 1
86 1
43 0
21 1
10 1
5 0
2 1
1 0
0 1
[1, 1, 0, 1, 1, 0, 1, 0, 1]
```

<!-- grade_id: cell-f96bf7f1f098f4a6, points: 0, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Unit test
# 0 Marks
assert(integer_digits(347,2) == [1, 0, 1, 0, 1, 1, 0, 1, 1])
print('Unit test passed! Run some of your own tests, making sure to try a variety of cases.')
```

Output:
```
173 1
86 1
43 0
21 1
10 1
5 0
2 1
1 0
0 1
AssertionError                            Traceback (most recent call last)
<ipython-input-49-d54949c6fe29> in <module>
      2 # Unit test
      3 # 0 Marks
----> 4 assert(integer_digits(347,2) == [1, 0, 1, 0, 1, 1, 0, 1, 1])
      5 print('Unit test passed! Run some of your own tests, making sure to try a variety of cases.')

AssertionError:
```

<!-- grade_id: cell-5922f55d871755cd, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2a(_globals)
```

Output:
```
Testing: arguments by position
3 0
1 1
0 1
integer_digits(6,2) returned the incorrect output
173 1
86 1
43 0
21 1
10 1
5 0
2 1
1 0
0 1
integer_digits(347,2) returned the incorrect output
34 7
3 4
0 3
integer_digits(347,10) returned the correct output

Testing: arguments by keyword
3 0
1 1
0 1
integer_digits(number=6,base=2) returned the incorrect output
173 1
86 1
43 0
21 1
10 1
5 0
2 1
1 0
0 1
integer_digits(base=2,number=347) returned the incorrect output
34 7
3 4
0 3
integer_digits(number=347,base=10) returned the correct output

Testing: edge cases
integer_digits(1,2) returned the incorrect output
integer_digits(base=7,number=1) returned the incorrect output
integer_digits(0,10) returned the incorrect output
<ipython-input-51-459f665e07b0> in <module>
      3 # Run this cell for grading
      4 _globals = globals()
----> 5 question2a(_globals)

~\OneDrive - Imperial College London\Introduction to Computation\Week 6\autograder_term2week3.py in question2a(_globals)
    175         re
```

(b) Adapt your <code>integer_digits</code> function so that the default value of <code>base</code> is 10. 

Test. If you've done it right, 

```python
integer_digits(347, 10)
integer_digits(347, 2)
integer_digits(number 347, base = 10)
```

and

```python
integer_digits(base = 10, number = 347)
```
should all still work, but in addition, 

```python
integer_digits(347)
```

and 

```python
integer_digits(number = 347)
```

should both return <code>[3, 4, 7]</code>.

<!-- grade_id: cell-b2b38ba00344d53d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a503a65776131ce2, points: 0, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Unit test
# 0 Marks
assert(integer_digits(347) == [3,4,7])
print('Unit test passed! Run some of your own tests, making sure to try a variety of cases.')
```

<!-- grade_id: cell-3a23b2de29ebc595, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2b(_globals)
```

(c) Adapt your code further so that <code>base</code> is only accepted as a keyword argument, but <code>number</code> can be positional or keyword.

If you've done it right, the following should work:

```python
integer_digits(347)
integer_digits(number = 347)
integer_digits(347, base = 2)
integer_digits(number = 347, base = 2)
integer_digits(base = 2, number = 347)
```

but 

```python
integer_digits(347, 2)
```

should no longer work.

Can you think of an advantage of this arrangement?

<!-- grade_id: cell-fda6ae637f58982c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-70f2ba6cc7659b0a, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2c(_globals)
```

## Question 3
(a) Write and test a hailstone sequence function which uses the infinite loop

```python
    while True:
```

instead of a loop with a terminating condition. This function should terminate at the first occurrence of 1. 

An <code>hailstone_step</code> function is provided for you below.

Note: autograding won't be able to tell whether you've used <code>while True</code>, but please do try.

<!-- grade_id: cell-74b97b68f8c335a0, points: 0, grade: false, solution: true -->
```python
def hailstone_step(a):
    if a % 2 == 0:
        return a//2
    else:
        return 3*a+1

def hailstone_sequence(a):
    if a==1:
        return [1]
    seq = [a]
    num = hailstone_step(a)
    seq.append(num)
    while True:
        if num == 1:
            return seq
        num = hailstone_step(num)
        seq.append(num)
        
print(hailstone_sequence(11))
print(hailstone_sequence(1))
```

Output:
```
[11, 34, 17, 52, 26, 13, 40, 20, 10, 5, 16, 8, 4, 2, 1]
[1]
```

<!-- grade_id: cell-1f24b08e62f943b4, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3a(_globals)
```

Output:
```
Testing: typical case
hailstone_sequence(11) returned the correct output

Testing: edge case
hailstone_sequence(1) returned the correct output

Question 3a: 2 out of 2
2
```

(b) Write and test an <code>aliquot_chain</code> function, called <code>aliquot_chain</code> which uses the infinite loop

```python
    while True:
```

instead of a loop with a terminating condition. It should terminate at the first occurrence of a number that has already been seen. You may want to reuse your <code>factor_sum</code> function from earlier in the module.

<!-- grade_id: cell-3d6fc02febf86f32, points: 0, grade: false, solution: true -->
```python
from math import sqrt, ceil

def factor_sum(a):
    sum = 0
    for r in range(1, ceil(sqrt(a))+1):
        if a%r == 0 and r != int(a/r):
            sum += r + int(a/r)
            print(r, a/r)
        elif r == sqrt(a):
            sum += r
            print(r, a/r)
    return sum

def aliquot_chain(a):
    alist = [a]
    while True:
        a = factor_sum(a) - a
        alist.append(a)
        if a in alist[0:-1]:
            return alist

aliquot_chain(30)
```

Output:
```
1 30.0
2 15.0
3 10.0
5 6.0
6 5.0
1 53.0
1 1.0
[30, 53, 1, 0, 0]
```

<!-- grade_id: cell-547043decd0b74d6, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3b(_globals)
```

Output:
```
Testing: typical case
1 282.0
2 141.0
3 94.0
6 47.0
1 294.0
2 147.0
3 98.0
6 49.0
7 42.0
14 21.0
1 390.0
2 195.0
3 130.0
5 78.0
6 65.0
10 39.0
13 30.0
15 26.0
1 618.0
2 309.0
3 206.0
6 103.0
1 630.0
2 315.0
3 210.0
5 126.0
6 105.0
7 90.0
9 70.0
10 63.0
14 45.0
15 42.0
18 35.0
21 30.0
1 1242.0
2 621.0
3 414.0
6 207.0
9 138.0
18 69.0
23 54.0
27 46.0
1 1638.0
2 819.0
3 546.0
6 273.0
7 234.0
9 182.0
13 126.0
14 117.0
18 91.0
21 78.0
26 63.0
39 42.0
1 2730.0
2 1365.0
3 910.0
5 546.0
6 455.0
7 390.0
10 273.0
13 210.0
14 195.0
15 182.0
21 130.0
26 105.0
30 91.0
35 78.0
39 70.0
42 65.0
1 5334.0
2 2667.0
3 1778.0
6 889.0
7 762.0
14 381.0
21 254.0
42 127.0
1 6954.0
2 3477.0
3 2318.0
6 1159.0
19 366.0
38 183.0
57 122.0
61 114.0
1 7926.0
2 3963.0
3 2642.0
6 1321.0
1 7938.0
2 3969.0
3 2646.0
6 1323.0
7 1134.0
9 882.0
14 567.0
18 441.0
21 378.0
27 294.0
42 189.0
49 162.0
54 147.0
63 126.0
81 98.0
1 12753.0
3 4251.0
9 1417.0
13 981.0
39 327.0
109 117.0
1 7267.0
13 559.0
43 169.0
1 785.0
5 157.0
1 163
```

(c) The following code is for a function called <code>power_of_two_divisor1</code>, which takes as its argument <code>n</code>, assumed to be a positive integer, and returns the highest value of <code>r</code> such that $2^r$ divides $n$; so for example

```python
power_of_two_divisor1(24)
```

returns 3, because $2^3$ divides 24.

```python
def power_of_two_divisor1(n):
    # initialize r
    r = 0
    
    # while loop, terminating when n becomes odd
    while n % 2 == 0:
        # increment n
        r += 1
        # divide n by 2
        n = n // 2
    
    # return final value of r
    return r
```

Test this function.

Then, write and test a version of it called <code>power_of_two_divisor2</code> that doesn't use a while loop at all, but instead uses the infinite loop

```python
for r in count()
```

An outline is provided below.

<!-- grade_id: cell-7ca2fe4de8b88de3, points: 0, grade: false, solution: true -->
```python
print(power_of_two_divisor1(112))



def power_of_two_divisor2(n):
    from itertools import count
    # infinite for loop
    for r in count():
        pass

```

Output:
```
4
```

<!-- grade_id: cell-b763bf1de9d34a56, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3c(_globals)
```

## Question 4

(a) Complete the following factorial function, <code>my_factorial</code>, which works using <b>recursion</b>: that is, by calling itself. Test your function.

<!-- grade_id: cell-8e24c72afd376c6e, points: 0, grade: false, solution: true -->
```python
def my_factorial(n):
    # base case
    if n==1 or n==0:
        return 1
    # recursion step
    else:
        return n * my_factorial(n-1)
my_factorial(10)
```

Output:
```
3628800
```

<!-- grade_id: cell-d8620df67910e0fb, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4a(_globals)
```

(b) Write and test a recursive version of your <code>hailstone_sequence</code> function, called <code>hailstone_sequence2</code>, which should work by calling itself instead of by a loop, and which should terminate at the first occurence of $1$. 

Autograding can't tell whether you've used recursion, but please do try!

<!-- grade_id: cell-483675ae72e172f2, points: 0, grade: false, solution: true -->
```python
def hailstone_step(a):
    if a % 2 == 0:
        return a//2
    else:
        return 3*a+1

def hailstone_sequence2(a):
    if a == 1 or a == 0:
        return lis.append(1)
    else:
        return
```

<!-- grade_id: cell-5b9dff1d06c22cb3, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4b(_globals)
```

(c) The following code defines an iterative function called <code>fibonacci</code>:

```python
def fibonacci(n):
    # initialize a and b
    a, b = 1, 1
    # the cases 1 and 2
    if n==1 or n==2:
        return 1
    else:
        # for loop
        for r in range(3,n+1):
            a, b = b, a + b
        return b
```

The following is an alternative, recursive version:

```python
def fibonacci2(n):
    # base case
    if n==1 or n==2:
        return 1
    # recursion step
    else:
        return fibonacci2(n-1) + fibonacci2(n-2)
```

Test that these two functions return the same arguments for $n=0, 1, 2, 3, \dots, 10$.

Autograding can't tell whether you've done that or not, but please do!

Calculate the 35th Fibonacci number using each of these two functions, then choose one of the following options:

1. Both functions give the same output for $n=35$, in about the same time.
2. Both functions give the same output for $n=35$, but <code>fibonacci</code> is the more efficient.
3. Both functions give the same output for $n=35$, but <code>fibonacci2</code> is the more efficient.
4. The two functions give different outputs for $n=35$, suggesting at least one of them has a bug.

<!-- grade_id: cell-fe4c35facb2ba712, points: 0, grade: false, solution: true -->
```python
question4c_answer = # (delete, and fill in the value 1, 2, 3 or 4)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1b5a4c306b8824ef, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question4c_answer == question4_c())
print('test case passed!')
```

(d) Write and test a recursive function called <code>fibonacci_recpair</code> which takes as its argument <code>n</code>, assumed to be an integer greater than 1, and returns, as a 2-tuple, the $(n-1)$th and $n$th Fibonacci numbers. So, for example, 

```python
fibonacci_recpair(10)
```

should return <code>(34, 55)</code>.

<!-- grade_id: cell-4e6a39f259575fb1, points: 0, grade: false, solution: true -->
```python
def fibonacci_recpair(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-dc17b901d6f3b7f6, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4d(_globals)
```

(e) Write and test a function called <code>fibonacci3</code> which, by calling <code>fibonacci_recpair</code>, calculates the $n$th Fibonacci number. (This won't work for $n=1$, so you'll need to make a special case of that.)

<!-- grade_id: cell-2b18346684193918, points: 0, grade: false, solution: true -->
```python
def fibonacci3(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e258528bc9ebe868, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4e(_globals)
```

(f) Test your code from part (e) on $n=35$, and perhaps on higher values, and then choose one of the following:

1. Recursion can throw up problems with execution time when a function calls itself more than once; these problems can't generally be solved.
2. Recursion can throw up problems with execution time when a function calls itself more than once, but it can be possible to find a clever work-around.
3. Recursion doesn't generally throw up problems with execution time at all.

<!-- grade_id: cell-0ce7f31ea1a61fcb, points: 0, grade: false, solution: true -->
```python
question4f_answer = # (delete, and fill in the value 1, 2, 3 or 4)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b0e5af91705dd496, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question4f_answer == question4_f())
print('test case passed!')
```

(g) Write and test a recursive implementation of your aliquot chain function, <code>aliquot_chain2</code>, which should terminate at the first occurrence of a number that has occurred before. You're likely to need your <code>factor_sum</code> function again.

Note: this is tricky. One thing to consider is writing, as with part (e), a recursive "inner" function and an "outer" function that simply calls it. In this case, the "inner" function would need to take as its argument not a single value but a list of values; in this way, it has access to information concerning what numbers have already occurred.

<!-- grade_id: cell-717274c920b3892e, points: 0, grade: false, solution: true -->
```python
def factor_sum(a):
    # YOUR CODE HERE
    raise NotImplementedError()

def aliquot_chain_inner(list_so_far):
    # YOUR CODE HERE
    raise NotImplementedError()

def aliquot_chain2(a):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-58522f35705ee5d9, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4g(_globals)
```

## Question 5 (Bonus Question)

The prime 41, can be written as the sum of six consecutive primes:
$$
41 = 2 + 3 + 5 + 7 + 11 + 13
$$
This is the longest sum of consecutive primes that adds to a prime below one-hundred. The longest sum of consecutive primes below one-thousand that adds to a prime contains 21 terms and is equal to 953. Write a function which takes in integer `n` and returns (i) the prime which corresponds to the longest sum of consecutive primes below `n`, and (ii) the length of this sum of primes. From our example above, if `n = 100`, then the function should return `(41, 6)`.

**Hint:** Use the sieve developed for primes from the earlier lab. 

**Note:** The time limits for the varying sized inputs.

<!-- grade_id: cell-746a820ef58a23b4, points: 0, grade: false, solution: true -->
```python
import numpy as np
def prime_chain(n):
    """
    Inputs: 
        n : the upper bound that the sequence of primes may not exceed - scalar
    Output:
        max_prime  : the maximum prime corresponding to the longest sum -scalar
        max_length : the length of the sequence of primes - scalar
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-677337a3ba85db4b, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# unit test
# 1 Mark

assert(41 == prime_chain(100)[0])
assert(6 == prime_chain(100)[1])
print('All Tests Passed!!!')

```

<!-- grade_id: cell-809ccc45c2022063, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# small input
# 1 Marks

assert(953 == prime_chain(1000)[0])
assert(21 == prime_chain(1000)[1])
print('All Tests Passed!!!')
```

<!-- grade_id: cell-18ec81e49e89b4ec, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# medium input
#--------------------------
# 15 second time limit
#--------------------------
# 1 Mark
import time
start      = time.time()
time_limit = 15

assert((9521, 65) == prime_chain(10000))
if abs(time.time() - start) < time_limit:
    print('All Tests Passed!!!')
else:
    raise AssertionError('Program Too Slow!!!')
 
```

<!-- grade_id: cell-d3223e3d75e6b1b6, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# large inputs
#--------------------------
# 30 second time limit
#--------------------------
# 1 Mark

start      = time.time()
time_limit = 30

assert((92951, 183) == prime_chain(100000))
if abs(time.time() - start) < time_limit:
    print('All Tests Passed!!!')
else:
    raise AssertionError('Program Too Slow!!!')

```

<!-- grade_id: cell-5257d126ea8742c1, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# large inputs
#--------------------------
# 60 second time limit
#--------------------------
# 1 Mark

start      = time.time()
time_limit = 60

assert((997651, 543) == prime_chain(1000000))
if abs(time.time() - start) < time_limit:
    print('All Tests Passed!!!')
else:
    raise AssertionError('Program Too Slow!!!')

```

<!-- grade_id: cell-41ef14b1d6c0a03e, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
#--------------------------
# 60 second time limit
#--------------------------
# huge inputs
# 2 Marks

start      = time.time()
time_limit = 60
assert((9951191, 1587) == prime_chain(10000000))
if abs(time.time() - start) < time_limit:
    print('All Tests Passed!!!')
else:
    raise AssertionError('Program Too Slow!!!')

```