```python
from autograder_week9 import *
_globals = globals()
```

# <center> Introduction to Mathematical Computing </center>
### <center>Phil Ramsden, Matt Woolway, Boumediene Hamzi</center>

# <center> Worksheet 4: Branching; iterable objects; comprehensions</center>

**NOTE:** When answering a question, remember to remove the `raise NotImplementedError()` line.

## Question 1

These questions are all about branching using <code>if</code>.

(a) Write a <code>for</code> loop that, starting with $a=36$, performs the following operation 30 times: if $a$ is even, make the new value of $a$ equal to $a/2$; otherwise, make it equal to $3a+1$. Your code should create a list of all the $a$-values, including the initial value, which you should call <code>alist</code>.

<!-- grade_id: cell-9fd6022d0f8e0c0f, points: 0, grade: false, solution: true -->
```python
# initialize a and alist
a = 36
alist = [a]

# for loop
for num in range(30):
    if a%2 == 0:
        a /= 2
    else:
        a = 3*a + 1
    alist.append(a)

# print final value of alist
print(alist)
```

<!-- grade_id: cell-eedcaa365a51095f, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1a(_globals)
```

(b) Once this sequence hits 1, we know what it's going to do, so there's really no point continuing. Write a new version, which instead of a <code>for</code> loop with a set number of iterations, uses a <code>while</code> loop that stops as soon as the value of $a$ becomes 1. The final value of <code>alist</code> should include 1.

<!-- grade_id: cell-d0812eeba6b267a9, points: 0, grade: false, solution: true -->
```python
# initialize a and alist
a = 36
alist = [a]

# while loop
while a != 1:
    if a%2 == 0:
        a /= 2
    else:
        a = 3*a + 1
    alist.append(a)


# print final value of alist
print(alist)
```

<!-- grade_id: cell-47bc7c560477788e, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1b(_globals)
```

c) The following is a piece of code that adds up all the factors of 24:

```python
# value of a
a = 24

# initialize factor sum
factor_sum = 0

# for loop
for r in range(1, a+1):
    # if r is a factor, add it to the factor sum
    if a % r == 0:
        factor_sum += r

# print the final value of the factor sum and the count variable
print(factor_sum)
print(r)
```

This is quite inefficient, though, because it uses 24 iterations where far fewer are needed. That's because once we know, say, 3 is a factor, we also know that 24/3 = 8 is a factor, so we might as well add that at the same time.

Write a new version of the code that uses this trick so that instead of 23, only about 4 iterations are needed.

<!-- grade_id: cell-e0d25c1168019f27, points: 0, grade: false, solution: true -->
```python
from math import sqrt, ceil

a=24
factor_sum=0

for r in range(1, ceil(sqrt(a))):
    if a % r == 0:
        factor_sum += r + (a/r)

print(factor_sum)
print(r)
```

<!-- grade_id: cell-4f9293c192941e67, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1c(_globals)
```

(d) Does your new code work if you change the value of $a$ to 36? Check the value it produces against the one produced by the old, inefficient code (which, though inefficient, is correct). If there is a bug, what's causing it?

Write a new version of your code from part (c), with any bugs removed; your new code should work both for $a=24$ and for $a=36$, and you should run the grading test with both values in turn. 

IMPORTANT: When you <em>submit</em>, make sure the value of `a` is set to 24 in question 1(c) and 36 here.

<!-- grade_id: cell-764d0f595f14ce23, points: 0, grade: false, solution: true -->
```python
a = 36
factor_sum = 0
for r in range(1, ceil(sqrt(a)+1)):
    if a % r == 0:
        print(r)
        print(a/r)
        factor_sum += r + (a/r)

factor_sum -= 6
# print the final value of the factor sum and the count variable
print(factor_sum)
print(r)
```

<!-- grade_id: cell-7f30e64e3dbda7b7, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1d(_globals)
```

## Question 2(i)

The function <code>ord</code> assigns a numerical code to each character; each character's code is unique:

```python
print(ord('A'))
print(ord('d'))
print(ord('?'))
```

(a) Using a single, one-line command, create a list consisting of all the character codes of the characters in the given string, <code>string1</code>. (Hint: use a <b>comprehension</b>.)

(b) Using a single, one-line command, create a list consisting of all the square roots of numbers in the list from part (a). (Remember, you may need to import a square root function.)

(c) Using the <code>zip</code> function, create a list, each of whose elements is a tuple of length 2. The first element of each tuple should be a value from the list from part (a), and the second should be the corresponding value from the list from part (b).

(d) Without using <code>zip</code>, but instead using a single-line comprehension over the list from part (a), create the same list of tuples as in part (c).

<!-- grade_id: cell-8e9d6b72facb2dc9, points: 0, grade: false, solution: true -->
```python
from math import sqrt

string1 = 'The quick brown fox jumps over the lazy dog'

s1list = [ord(char) for char in string1]
s1list2 = [sqrt(item) for item in s1list]
zip1 = list(zip(s1list, s1list2))
zip2 = [(s1list[i], s1list2[i]) for i in range(len(s1list))]


q2ia_answer = s1list
q2ib_answer = s1list2
q2ic_answer = zip1
q2id_answer = zip2

print(q2ia_answer)
print(q2ib_answer)
print(q2ic_answer)
print(q2id_answer)
```

<!-- grade_id: cell-dd0ddb84d57eaced, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2i(_globals)
```

## Question 2(ii)

The operator <code>in</code> can be used to check whether a particular character belongs to a particular string, or whether a particular element belongs to a particular list:

```python
no_e_string = 'This longish string, I think possibly slightly oddly, is missing a particular symbol, which ordinarily you would bank on occurring.'

print(no_e_string)
print('e' in no_e_string)
print('a' in no_e_string)
```

```python
odd_list = list(range(1,100,2))

print(odd_list)
print(64 in odd_list)
print(65 in odd_list)
```

(a) Using a single-line comprehension, find out exactly which <i>lower-case</i> letters of the alphabet are in <code>no_e_string</code>. Do this by creating a list of 26 2-tuples; the first element of each tuple should be a lower-case letter of the alphabet, and the second should either be <code>True</code> or <code>False</code>. 

(b) Using a single-line comprehension based on the <code>pow</code> function, create a list of the residues of $2^n$ modulo $95$ for $n=0, 1, 2, \dots, 94$.

(c) Find out which of the integers between $0$ and $94$ inclusive belong to your list from part (b). Do this by creating a list of 95 lists of length 2; the first element should be an integer between $0$ and $94$, and the second should either be <code>True</code> or <code>False</code>.

<!-- grade_id: cell-3a14546de16db48d, points: 0, grade: false, solution: true -->
```python
no_e_string = 'This longish string, I think possibly slightly oddly, is missing a particular symbol, which ordinarily you would bank on occurring.'
alphabet = 'abcdefghijklmnopqrstuvwxyz'

alist = [(alphabet[i], alphabet[i] in no_e_string) for i in range(len(alphabet))]
powlist = [pow(2, n, 95) for n in range(95)]
powlist2 = [(n , n in powlist) for n in range(95)]

q2iia_answer = alist
q2iib_answer = powlist
q2iic_answer = powlist2

print(q2iia_answer)
print(q2iib_answer)
print(q2iic_answer)
```

<!-- grade_id: cell-53cf9a2d252b9b17, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2ii(_globals)
```

## Question 3

Suppose Alice wants to send a message to Bob, but their enemy Chuck is listening in. They might want to use a "secret code"; technically a <b>cipher</b>.

One of the simplest kinds is called a <b>Caesar cipher</b>. This works by shifting every character along by the same amount: for example, A goes to L, B goes to M, C goes to N, etc.

It's quite good to work not with a 26-letter alphabet but with a fairly full set of keyboard characters. There are about 95 of these, with character codes between 32 and 126:

```python
"".join([chr(x) for x in range(32,95)])
```

Then a Caesar cipher might consist of an encryption rule along the lines of

"Shift each character along by 37, wrapping round cyclically when you get to 95."

If we let each character in the original message (the <b>plaintext</b>) be represented by a number $p$ between 0 and 94, this rule can be written

$$c = (p + 37) \mod 95,$$

where $c$ is the corresponding number in the encrypted message (the <b>ciphertext</b>). 

A complete encryption in Python might look like this:

```python
# The original message
plaintext = 'My name is Bond, James Bond.'

# The encryption key
n = 95
e = 37

# Replacing every character with a number between 0 and 94
plaintext_numbers = [ord(x) - 32 for x in plaintext]

# Shifting along by 37, wrapping round at 95
ciphertext_numbers = [(x + e) % n for x in plaintext_numbers]

# Converting back into characters
ciphertext = "".join([chr(x+32) for x in ciphertext_numbers])

print(plaintext_numbers)
print(ciphertext_numbers)
print(ciphertext)
```

That's what Alice does. Bob's problem is then to recover the original message: to <b>decrypt</b> the ciphertext. Because this is modular arithmetic, he can undo this addition of a positive integer by addition of another positive integer!

## 3 (i)

Play the role of Bob. 

(a) Starting with the ciphertext 'r?E4'3+E/9Eg54\*QEo'3+9Eg54\*S', begin by converting this into a list of integers between 0 and 94.

(b) Now, undo the addition of 37, modulo 95, by addition of another positive integer. 

(c) Finally, convert this list of integers into a string of characters.

<!-- grade_id: cell-2a3d027244b05804, points: 0, grade: false, solution: true -->
```python
# Slightly complicated code; this deletes all previous definitions, but only if they exist
# You needn't try to understand this yet
try:
    del plaintext, plaintext_numbers, ciphertext_numbers, ciphertext
except NameError:
    pass

ciphertext = "r?E4'3+E/9Eg54*QEo'3+9Eg54*S"

# Part (a)
ciphertext_numbers = [ord(x)-32 for x in ciphertext]

# Part (b)
# decryption key; fill in the correct values, making sure d is positive
n = 95
d = 58

decrypt_numbers = [(x + d) % n for x in ciphertext_numbers]

# Part (c)
decrypt = "".join([chr(num+32) for num in decrypt_numbers])

print(ciphertext_numbers)
print(decrypt_numbers)
print(decrypt)
```

<!-- grade_id: cell-5a29c97b2d547b00, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3i(_globals)
```

The Caesar cipher is based on modular addition. It's possible to base a cipher on modular multiplication instead, using an algorithm such as 

$$c = p * 37 \mod 95.$$

## 3 (ii)
(a) Play the role of Alice: encrypt the message 'Good evening, Mr Bond, I've been expecting you.' using the above linear cipher.

<!-- grade_id: cell-7c83b77d2c6cdcf1, points: 0, grade: false, solution: true -->
```python
plaintext = "Good evening, Mr Bond, I've been expecting you."

# The encryption key
n = 95
e = 37

# Replacing every character with a number between 0 and 94
plaintext_numbers = [ord(char)-32 for char in plaintext]

# Multiplying by 37, wrapping round at 95
ciphertext_numbers = [(num*e)%n for num in plaintext_numbers]

# Converting back into characters
ciphertext = "".join([chr(num+32) for num in ciphertext_numbers])

print(plaintext_numbers)
print(ciphertext_numbers)
print(ciphertext)
```

<!-- grade_id: cell-9a79b700b170b800, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3iia(_globals)
```

Now to play the role of Bob. This is a little tricky, because you first have to work out what Bob's decryption algorithm ought to be. 

It turns out that we can undo multiplication by $e$, modulo 95, using multiplication by another number $d$, modulo 95. This number has the property that the residue of $ed$, modulo 95, is 1.

(b) Generate a list of 94 2-tuples; the first element of each tuple should be a number $x$ between 0 and 94, and the second should be the residue of $37*d$, modulo 95.

(c) By eye, find the tuple whose second element is 1, and hence find the decryption key, $d$.

(d) Decrypt Alice's message using multiplication by $d$, modulo 95.

<!-- grade_id: cell-0c134f764924d2a4, points: 0, grade: false, solution: true -->
```python
# redefine e and n in case they've been cleared
e = 37
n = 95

# Part (b)
residue_tuples = [(x, (x*37)%95) for x in range(95)]
print(residue_tuples)

# Part (c)
d = 18
print(d)

# Part (d)
# Slightly complicated code; this deletes all previous definitions, but only if they exist
# You needn't try to understand this yet
try:
    del plaintext, plaintext_numbers, ciphertext_numbers, ciphertext
except NameError:
    pass

# the ciphertext
ciphertext = '2iiN sOsDID^` Ry 7iDN` |eOs cssD s:/s)dID^ _i*K'
# convert into a list of integers between 0 and 94
ciphertext_numbers = [ord(char)-32 for char in ciphertext]
# decrypt, using multiplication by d, modulo n (where n is 95)
decrypt_numbers = [(x*d)%n for x in ciphertext_numbers]
# convert into a string
decrypt = "".join([chr(x+32) for x in decrypt_numbers])

print(ciphertext_numbers)
print(decrypt_numbers)
print(decrypt)
```

<!-- grade_id: cell-03fba32d87f9f597, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3iibcd(_globals)
```

(e) What would have happened if Alice had used $e=38$ instead of $e=37$? Select one of the following:

<ol>
    <li>Bob could have decrypted the message using the same decryption key.</li>
    <li>Bob could have decrypted the message, but would need to calculate a new decryption key.</li>
    <li>Bob could not have reliably decrypted the message, because there would have been ambiguities.</li>
</ol>

<!-- grade_id: cell-0a1a90d175332d71, points: 0, grade: false, solution: true -->
```python
q3e_answer = 3
```

<!-- grade_id: cell-a3a2773c75b62ad2, points: 1, grade: true, solution: false -->
```python
#1 Mark
assert(q3e_answer == question3e())
print("Question 3e passed!")
```

## 3 (iii)

(a) Encrypt the message 

"Vodka Martini, shaken not stirred." 

using the algorithm

$$c = p^{13}\mod 95.$$

(Don't forget about the <code>pow</code> function.)

(b) Now to find the decryption key. It turns out that because 95 is the product of the primes 5 and 19, we want as our $d$ a number with the property that the residue of $d*e$ modulo 72 is 1. (Why 72? It has to do with the fact that $72 = (5 - 1) \times (19 - 1)$; more about this in one of your other modules.)

Find a number between 0 and 71 with this property, and assign its value to the variable $d$.

(c) Finally, decrypt the message using the decryption algorithm

$$p = c^d \mod n.$$

<!-- grade_id: cell-256adea0bb7003ae, points: 0, grade: false, solution: true -->
```python
# Part(a)

# Fill in your own code here, ending with...

plaintext = 'Vodka Martini, shaken not stirred'
p_numbers = [ord(char)-32 for char in plaintext]
c_numbers = [pow(x, 13, 95) for x in p_numbers]
ciphertext = "".join(chr(x+32) for x in c_numbers)

print(ciphertext)

# Part (b)

decrypt_keys = [(x, (x*13)%72) for x in range(72)]
print(decrypt_keys)

d = 61

print(d)

# Part (c)
c_numbers2 = [ord(char)-32 for char in ciphertext]
decrypt_numbers = [pow(x, d, 95) for x in c_numbers2]
decrypt = "".join([chr(x+32) for x in decrypt_numbers])

print(decrypt)
```

<!-- grade_id: cell-6bf2533d39eb367a, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3iii(_globals)
```

## Question 4

(a) First, create four objects called <code>myrange</code>, <code>mylist</code>, <code>mytuple</code> and <code>myslice</code>, which should be defined as follows.

<ul>
    <li><code>myrange</code> should be a range, containing all the integers between 97 and 102 inclusive.</li>
    <li><code>mylist</code> should be a list, again containing all the integers between 97 and 102 inclusive.</li>
    <li><code>mytuple</code> should be a tuple, again containing all the integers between 97 and 102 inclusive.</li>
    <li><code>myslice</code> should be an object created using the <code>count</code> and <code>islice</code> functions from the <code>itertools</code> module, again containing all the integers between 97 and 102 inclusive.</li>
</ul>

<!-- grade_id: cell-c4c91cf440f21ed4, points: 0, grade: false, solution: true -->
```python
from itertools import count, islice

myrange = range(97, 103)
mylist = [x for x in myrange]
mytuple = tuple(mylist)
myslice = islice(count(97, 1), 6)

print(myrange)
print(mylist)
print(mytuple)
print(myslice)
```

<!-- grade_id: cell-584f73de5583f23f, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4a(_globals)
```

(b) Run the following code:

```python
print([n**2 for n in myrange])
print(list(myrange))
```

Repeat for the other three objects, and then choose one of the following:

<ol>
    <li>Running a comprehension on all four objects works, and leaves each object unchanged.</li>
    <li>Running a comprehension on all four objects works, and leaves three of the objects unchanged while changing the other one.</li>
    <li>Running a comprehension on all four objects works, and leaves two of the objects unchanged while changing the other two.</li>
    <li>Running a comprehension on all three objects works, and leaves one of the objects unchanged while changing the other three.</li>
    <li>Running a comprehension does not work on all of the objects.</li>
</ol>

<!-- grade_id: cell-54054268baec99be, points: 0, grade: false, solution: true -->
```python
print([n**2 for n in mylist])
print(mylist)
print([n**2 for n in mytuple])
print(mytuple)
print([n**2 for n in myslice])
print(myslice)

q4b_answer = 2

```

<!-- grade_id: cell-a1d96c5ad00f6a61, points: 1, grade: true, solution: false -->
```python
#1 Mark
assert(q4b_answer == question4b())
print("Question 4b passed!")
```

(c) In the video lectures, we checked that ranges, lists, tuples and islice objects all belonged to the class Iterable, and that the first three belonged to the class Sequence. Find out which of them belong to the class Iterator (this class, like the other two, can be found in the module collections.abc).

Give your answers enclosed in curly brackets; for example, if you think ranges and tuples belong to the class Iterator, then your answer would be {1, 3}.

<!-- grade_id: cell-feb9b634c72d96eb, points: 0, grade: false, solution: true -->
```python
q4c_answer = # enter your selections in curly brackets
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a14716840a76e732, points: 1, grade: true, solution: false -->
```python
#1 Mark
assert(q4c_answer == question4c())
print("Question 4c passed!")
```

## Question 5

(a) Create a list called <code>ints</code>, consisting of all the integers between 2 and 400 inclusive. (Hint: use the <code>range</code> function, and wrap your <code>range</code> object in <code>list</code>.)

<!-- grade_id: cell-abd9e74afe06dfd9, points: 0, grade: false, solution: true -->
```python
ints = list(range(2, 401))
ints_old = list(range(2, 401))
```

<!-- grade_id: cell-e446afec6056fc52, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question5a(_globals)
```

(b) Run the following code:

```python
a = 2
ints = [n for n in ints if n == a or n % a > 0]
print(ints)
```

Which of the following is true?

<ol>
    <li>This code leaves the contents of the list <code>ints</code> unchanged.</li>
    <li>This code removes all the even numbers from <code>ints</code>.</li>
    <li>This code removes all the odd numbers from <code>ints</code>.</li>
    <li>This code removes <em>nearly</em> all the even numbers from <code>ints</code>.</li>
    <li>This code removes <em>nearly</em> all the odd numbers from <code>ints</code>.</li>
</ol>

<!-- grade_id: cell-a7626d40c088979e, points: 0, grade: false, solution: true -->
```python
q5b_answer = 4
```

<!-- grade_id: cell-e30a08bf39bd9c38, points: 1, grade: true, solution: false -->
```python
#1 Mark
assert(q5b_answer == question5b())
print("Question 5b passed!")
```

(c) Repeat, this time setting <code>a = 3</code>. Then again, using <code>a = 5</code>. Carry on until the contents of the list <code>ints</code> no longer change.

<!-- grade_id: cell-87aab97af2a51212, points: 0, grade: false, solution: true -->
```python
a = 3
while ints_old != ints:
    ints_old = ints
    ints = [n for n in ints if n==a or n % a > 0]
    a += 2
```

<!-- grade_id: cell-6bbdc5c3f5ac35c7, points: 0, grade: false, solution: true -->
```python
print(ints)
```

(d) What single word in English describes the numbers that form the contents of <code>ints</code> at the end of this process? They are all... what?

<!-- grade_id: cell-5ca5d9f319bc4f1b, points: 0, grade: false, solution: true -->
```python
q5d_answer = 'primes'
```

<!-- grade_id: cell-04e772337ae73e99, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question5d(_globals)
```

## Question 6 (Bonus Question)

The following question is a bonus question, and is for enrichment only. Please do not worry if it seems challenging for now; as it assumes knowledge of functions. You can always come back to it at a later stage in the course.



Reconsider **Question 1**, recall that the following iterative sequence is defined for the set of positive integers:
$$
\begin{cases}
a \rightarrow a/2\ \ (n\ \ \text{is even}) \\
a \rightarrow 3a + 1\ \ (n\ \ \text{is odd})
\end{cases}
$$
Using this rule and starting with 13, we generate the following sequence:
$$
13 \rightarrow 40 \rightarrow 20 \rightarrow 10 \rightarrow 5 \rightarrow 16 \rightarrow 8 \rightarrow 4 \rightarrow 2 \rightarrow 1
$$
It can be seen that this sequence (starting at 13 and finishing at 1) contains 10 terms. Although it has not been proven, it is thought that all starting numbers finish at 1. 
Write a function `longest_sequence` which takes in a upper bound integer `start` and finds which integer less than or equal to `start` that produces the longest sequence. The function should return two outputs, namely the length of the longest sequence and the integer which yields it.

<!-- grade_id: cell-e7c5cf586b8cdd14, points: 0, grade: false, solution: true -->
```python
def gen_seq_count(first):
    if not isinstance(first, int):
        return "no"
    count = 0
    print(first)
    while first != 1:
        if first%2 == 0:
            first /= 2
        else:
            first = 3*first + 1
        count += 1
        print(first)
    return count+1

def longest_sequence(start):
    maxlen = 0
    maxnum = 0
    for num in range(1, start+1):
        len1 = gen_seq_count(num)
        if len1 > maxlen:
            maxlen = len1
            maxnum = num
    return maxlen, maxnum

longest_sequence(3)
longest_sequence(2)
```

<!-- grade_id: cell-3c77a49e00ded13e, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# unit tests
n1 = longest_sequence(10)
n2 = longest_sequence(1000)
question6(n1, n2)
```

<!-- grade_id: cell-74d5278dbfc0e7d0, points: 2, grade: true, solution: false -->
```python
# Run this cell to see if your implementation is fast enough
# 2 Marks
# 15 second time limit
question6_timed(_globals)
```

```python

```