```python
from autograder_week10 import *
_globals = globals()
```

Output:
```
Autograder loaded successfully!
Remember to always restart and run all from the Kernel menu before submitting!
```

# <center> Introduction to Mathematical Computing </center>
### <center>Phil Ramsden, Matt Woolway, Boumediene Hamzi</center>

# <center> Worksheet 5: Functions</center>

**NOTE:** When answering a question, remember to remove the `raise NotImplementedError()` line.

## Question 1

(a) Write and test a function called `xsqdplusone` that takes one argument $x$ (which you can assume to be numerical) and returns the value of $x^2 + 1$. Your function should return a number of the same type as $x$.

<!-- grade_id: cell-b00091df7b05694f, points: 0, grade: false, solution: true -->
```python
def xsqdplusone(x):
    return x**2 + 1
```

<!-- grade_id: cell-35116d17e85ec49c, points: 0, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Unit test
# 0 Marks
assert(xsqdplusone(3+2j) == (6+12j))
print('Unit test passed! Run some of your own tests, making sure to try a variety of cases.')
```

Output:
```
Unit test passed! Run some of your own tests, making sure to try a variety of cases.
```

<!-- grade_id: cell-9f3ccf3b0dcc3a1a, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1a(_globals)
```

Output:
```
Your function returned the correct output for the input 1
Your function returned the correct output for the input 0.357
Your function returned the correct output for the input (3+2j)
Question 1a passed!
1
```

(b) Write and test a function called <code>hailstone_step</code> that takes one argument $a$ (which you can assume to be a positive integer) and returns an integer equal to $a/2$ if $a$ is even, and $3a+1$ if $a$ is odd.

<!-- grade_id: cell-7582dbb0352598f5, points: 0, grade: false, solution: true -->
```python
def hailstone_step(a):
    if a%2 == 0:
        return int(a/2)
    else:
        return 3*a + 1

print(hailstone_step(28))
```

Output:
```
14
```

<!-- grade_id: cell-5d58e3855945bf93, points: 0, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Unit test
# 0 Marks
assert(hailstone_step(7) == 22)
print('Unit test passed! Run some of your own tests, making sure to try a variety of cases.')
```

Output:
```
Unit test passed! Run some of your own tests, making sure to try a variety of cases.
```

<!-- grade_id: cell-03a43bd44d34aae1, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1b(_globals)
```

Output:
```
Test 1: positive odd argument
Your function returned numerically correct output for the input 1
Your function returned output of the correct numerical type for the input 1
Test 1 passed!

Test 2: positive even argument
Your function returned numerically correct output for the input 28
Your function returned output of the correct numerical type for the input 28
Test 2 passed!

2 marks out of 2
2
```

(c) Write a function called <code>my_factorial</code> that takes one argument $n$ (which you can assume to be a non-negative integer) and returns its factorial, $n!$.

<!-- grade_id: cell-72ce8a0b51ff800f, points: 0, grade: false, solution: true -->
```python
def my_factorial(n):
    fac = 1
    for num in range(1, n+1):
        fac *= num
    return fac
```

<!-- grade_id: cell-3bb5ea38caa35539, points: 0, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Unit test
# 0 Marks
assert(my_factorial(5) == 120)
print('Unit test passed! Run some of your own tests, making sure to try a variety of cases, including edge cases.')
```

Output:
```
Unit test passed! Run some of your own tests, making sure to try a variety of cases, including edge cases.
```

<!-- grade_id: cell-12fd18ed347a5d4e, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1c(_globals)
```

(d) Write and test a function called <code>fibonacci</code>, which takes as its argument $n$, which you can assume to be a positive integer, and returns the $n$th Fibonacci number. (Note: the first Fibonacci number can be taken to be 1, the second 1, the third 2, etc, so
```python
fibonacci(5)
```
should return 5.)

<!-- grade_id: cell-083814268f826c5d, points: 0, grade: false, solution: true -->
```python
def fibonacci(n):
    if n==1 or n==2:
        return 1
    else:
        return fibonacci(n-1) + fibonacci(n-2)
    
print(fibonacci(3))
```

Output:
```
2
```

<!-- grade_id: cell-0e433bf224e67d1a, points: 0, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Unit test
# 0 Marks
assert(fibonacci(5) == 5)
print('Unit test passed! Run some of your own tests, making sure to try a variety of cases.')
```

Output:
```
Unit test passed! Run some of your own tests, making sure to try a variety of cases.
```

<!-- grade_id: cell-512e72d91dabe64f, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1d(_globals)
```

Output:
```
Test 1: a list of typical cases, tested with a comprehension
Your function returned the correct output for all of the integers 3, 4, ..., 10
Test 1 passed!

Test 2: edge cases
Your function returned the correct output for the input 1
Your function returned the correct output for the input 2
Test 2 passed!

2 marks out of 2
2
```

(e) Write and test a function called <code>factor_sum</code> that takes as its argument <code>a</code>, which you can assume to be a positive integer, and returns the sum of all its factors including itself. Try to make your function reasonably efficient, without too many unnecessary iterations.

<!-- grade_id: cell-ed953a434d1a00e3, points: 0, grade: false, solution: true -->
```python
from math import sqrt, ceil

def factor_sum(a):
    sum = 0
    for r in range(1, ceil(sqrt(a))+1):
        if a%r == 0 and r != int(a/r):
            sum += r + int(a/r)
            print(r, a/r)
        elif r == sqrt(a):
            sum += r
            print(r, a/r)
    return sum

print(factor_sum(15))
print(factor_sum(36))
```

Output:
```
1 15.0
3 5.0
24
1 36.0
2 18.0
3 12.0
4 9.0
6 6.0
91
```

<!-- grade_id: cell-9124bb3900fcaf38, points: 0, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Unit test
# 0 Marks
assert(factor_sum(24) == 60)
print('Unit test passed! Run some of your own tests, making sure to try a variety of cases.')
```

Output:
```
1 24.0
2 12.0
3 8.0
4 6.0
Unit test passed! Run some of your own tests, making sure to try a variety of cases.
```

<!-- grade_id: cell-b19e110c9aa2d269, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1e(_globals)
```

Output:
```
Test 1: typical case, non-square
1 24.0
2 12.0
3 8.0
4 6.0
Your function returned the correct output for the input 24
Test 1 passed!

Test 2: typical case, square
1 36.0
2 18.0
3 12.0
4 9.0
6 6.0
Your function returned the correct output for the input 36
Test 2 passed!

Test 3: edge cases
1 1.0
Your function returned the correct output for the input 1
Your function returned the correct output for the input 0
Test 3 passed!

Test 4: efficiency
1 1000000.0
2 500000.0
4 250000.0
5 200000.0
8 125000.0
10 100000.0
16 62500.0
20 50000.0
25 40000.0
32 31250.0
40 25000.0
50 20000.0
64 15625.0
80 12500.0
100 10000.0
125 8000.0
160 6250.0
200 5000.0
250 4000.0
320 3125.0
400 2500.0
500 2000.0
625 1600.0
800 1250.0
1000 1000.0
Your function seemed to calculate the factor sum of one million in good time.
Test 4 passed!

4 marks out of 4
4
```

(f) Write a function called <code>nest</code>, which takes as its arguments a function <code>f</code>, a starting value <code>x0</code> and a number of iterations <code>n</code>, and returns 

<center>
    <code>f(f(f(...f(x0)...))),</code>
</center>

the <code>n</code>th iterate of <code>f</code> starting at <code>x0</code>.

Test your function. The following code should return the value 

$$\cos(\cos(\cos(\cos(\cos(1.0))))) = 0.7013687736227566,$$ 

for example (though you should test more than once).

```python
from math import cos
nest(cos, 1.0, 5)
```

Using your function, calculate the 20th iterate of $f$ starting at $0.3$, where

$$f(x) = 4\,x\,(1-x).$$

Make sure to pass this to the variable <code>x20</code>.

(Hint; for the last part, you may wish to use something called a <b>lambda-expression</b>; take a sneak peek at the notes for Term 2, Week 1. A lambda-expression is not required for this question, but you might find it convenient.)

<!-- grade_id: cell-f036c883340009fc, points: 0, grade: false, solution: true -->
```python
from math import cos

def nest(f, x0, n):
    x = x0
    for i in range(n):
        x = f(x)
    return x
    

x20 = nest(lambda n: 4*n*(1-n), 0.3, 20)
print(x20)
```

Output:
```
0.9417846056085262
```

<!-- grade_id: cell-eb91dac56e3e4ea7, points: 0, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Unit test
# 0 Marks
from math import cos
from numpy import isclose
assert(isclose(nest(cos, 1.0, 5), 0.7013687736227566))
print('Unit test passed! Run some of your own tests, making sure to try a variety of cases.')
```

Output:
```
Unit test passed! Run some of your own tests, making sure to try a variety of cases.
```

<!-- grade_id: cell-8dda4608d9bf0aec, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1f(_globals)
```

Output:
```
Test 1
nest(cos, 1.0, 5) returned the correct value
Test 1 passed!

Test 2
An attempt to iterate the function that maps x to exp(-x) 10 times with x0 = 1.0 returned the correct value
Test 2 passed!

Test 3
An attempt to iterate the function that maps x to 4*x*(1-x) 20 times with x0 = 0.3 returned the correct value
Test 3 passed!

Test 4: the value of x20
Your calculated value of x20 is correct
Test 4 passed!

4 marks out of 4
4
```

## Question 2


(a) Write and test a function called <code>hailstone_sequence</code> which takes as its arguments a starting value <code>a</code> and a number of iterations <code>n</code>, both assumed to be positive integers. It should then return, as a list, the sequence generated by iterating "if <code>a</code> is even, halve it; otherwise, multiply it by 3 and add 1" <code>n</code> times; the list should be of length <code>n+1</code>.

(Hint: you may wish to call the function you've already written called <code>hailstone_step</code>.)

<!-- grade_id: cell-bf69a919a8656af0, points: 0, grade: false, solution: true -->
```python
def hailstone_sequence(a, n):
    arr = [a]
    num = a
    for i in range(n):
        num = hailstone_step(num)
        arr.append(num)
    return arr
```

<!-- grade_id: cell-bad32cb1372c0531, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2a(_globals)
```

Output:
```
hailstone_sequence(36, 20) returned a list
Your function returns a sequence of ints
hailstone_sequence(36, 20) returned a sequence of the correct length
hailstone_sequence(36, 20) returned the correct value
4 marks out of 4
4
```

(b) Write and test a function called <code>fibonacci_sequence</code> which takes as its argument $n$ (which you can assume to be a positive integer) and returns a list consisting of the Fibonacci sequence up to the $n$th term. The first Fibonacci number should be treated as 1, and the second as 1, so that
```python
fibonacci_sequence(5)
```
should return
```
[1, 1, 2, 3, 5]
```

<!-- grade_id: cell-819fd6713dfb2bbd, points: 0, grade: false, solution: true -->
```python
def fibonacci_sequence(num): 
    arr = []
    for i in range(1, num+1):
        arr.append(fibonacci(i))
    return arr

print(fibonacci_sequence(3))
```

Output:
```
[1, 1, 2]
```

<!-- grade_id: cell-3e5dee79d8ab2138, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2b(_globals)
```

Output:
```
Test 1: typical case
fibonacci_sequence(10) returned the correct value
Test 1 passed!

Test 2: edge cases
fibonacci_sequence(2) returned the correct value
fibonacci_sequence(1) returned the correct value
Test 2 passed!

2 marks out of 2
2
```

(c) Write a function called <code>aliquot_chain</code> which takes as its argument a starting value <code>a</code> and a number of iterations <code>n</code>, both non-negative integers. It should return a list of length <code>n+1</code>, starting with <code>a</code>, where each number in the list (starting from the second) is the sum of the factors of the previous number <em>not including that number itself</em>.

Test your function: the code

<code>
aliquot_chain(24, 5)
</code>

should return the value <code>[24, 36, 55, 17, 1, 0]</code>, for example.

<!-- grade_id: cell-39b4848ed884388d, points: 0, grade: false, solution: true -->
```python
from math import ceil, sqrt

def factor_sum2(num):
    fac_sum = 0
    for r in range(2, ceil(sqrt(num))+1):
        if num%r == 0 and r != int(num/r):
            fac_sum += r + int(num/r)
            print(r, num/r)
        elif r == sqrt(num):
            fac_sum += r
            print(r, num/r)
    return fac_sum+1


def aliquot_chain(a, n):
    num = a
    arr = [a]
    for i in range(n):
        num = factor_sum2(num)
        if i==n-1 and num == 1:
            arr.append(num-1)
        else:
            arr.append(num)
    return arr

print(aliquot_chain(1, 5))
```

Output:
```
[1, 1, 1, 1, 1, 0]
```

<!-- grade_id: cell-4eaaee03a4b30bee, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2c(_globals)
```

Output:
```
Test 1: typical case
2 12.0
3 8.0
4 6.0
2 18.0
3 12.0
4 9.0
6 6.0
5 11.0
aliquot_chain(24, 5) returned the correct value
Test 1 passed!

Test 2: edge cases
2 12.0
3 8.0
4 6.0
aliquot_chain(24, 1) returned the correct value
aliquot_chain(24, 0) returned the correct value
aliquot_chain(1, 5) returned an incorrect value
aliquot_chain(0, 5) returned an incorrect value
Test 2 failed!

1 mark out of 2
1
```

(d) Write and test a function called <code>nest_list</code> which takes as its arguments a function <code>f</code>, a starting value <code>x0</code> and a number of iterations <code>n</code>, and returns the list

<center>
    <code>[x0, f(x0), f(f(x0)), ..., f(f(f(...f(x0)...)))],</code>
</center>

(that is, all <code>n+1</code> iterates of <code>f</code>, from the starting value <code>x0</code> to the <code>n</code>th).

Use your function to calculate the terms in the sequence generated by 20 iterations of $f$, where

$$f(x) = 4\,x\,(1-x),$$

starting from $x=0.3$. Pass this to the variable <code>xlist20</code>.

<!-- grade_id: cell-2ae8066aa476fbc7, points: 0, grade: false, solution: true -->
```python
def nest_list(f, x0, n):
    num = x0
    arr = [x0]
    for i in range(n):
        num = f(num)
        arr.append(num)
    return arr

xlist20 = nest_list(lambda x: 4*x*(1-x), 0.3, 20)
print(xlist20)
```

Output:
```
[0.3, 0.84, 0.5376000000000001, 0.9943449599999999, 0.02249224209039382, 0.08794536454456375, 0.32084390959875014, 0.8716123810885569, 0.4476169528867727, 0.9890240655005337, 0.043421853445318986, 0.1661455843547689, 0.5541649166167251, 0.9882646472316129, 0.04639053705515447, 0.17695382050755523, 0.5825646636613406, 0.9727323052579588, 0.10609667026198408, 0.3793606672852156, 0.9417846056085262]
```

<!-- grade_id: cell-6ecd8661ad6362cf, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2d(_globals)
```

Output:
```
Test 1
nest_list(cos, 1.0, 5) returned the correct value
Test 1 passed!

Test 2
An attempt to iterate the function that maps x to exp(-x) 10 times with x0 = 1.0 returned the correct value
Test 2 passed!

Test 3
An attempt to iterate the function that maps x to 4*x*(1-x) 20 times with x0 = 0.3 returned the correct value
Test 3 passed!

Test 4: the value of xlist20
Your calculated value of x20 is correct
Test 4 passed!

4 marks out of 4
4
```

## Question 3

(a) Write and test a version of your hailstone sequence function, called <code>hailstone_sequence2</code>, which takes only the argument <code>a</code> and which terminates not after a set number of iterations but when the sequence hits the value 1. (You probably want to use a <code>while</code> loop.)

<!-- grade_id: cell-48e58cfbd11e3e3b, points: 0, grade: false, solution: true -->
```python
def hailstone_sequence2(a):
    num = a
    arr = [a]
    while num != 1:
        if num%2 == 0:
            num /= 2
            num = int(num)
        else:
            num = 3*num + 1
        arr.append(num)
    return arr

print(hailstone_sequence2(24))
            
```

Output:
```
[24, 12, 6, 3, 10, 5, 16, 8, 4, 2, 1]
```

<!-- grade_id: cell-1c2b062449dc979c, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3a(_globals)
```

Output:
```
hailstone_sequence2(30) returned a list
Your function returns a sequence of ints
hailstone_sequence2(36, 20) returned a sequence terminating in 1
hailstone_sequence2(36) returned the correct value
hailstone_sequence2(1) returned the correct value
5 marks out of 5
5
```

<!-- grade_id: cell-760a9e06454837a1, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3a(_globals)
```

Output:
```
hailstone_sequence2(30) returned a list
Your function returns a sequence of ints
hailstone_sequence2(36, 20) returned a sequence terminating in 1
hailstone_sequence2(36) returned the correct value
hailstone_sequence2(1) returned the correct value
5 marks out of 5
5
```

(b) Write a version of your aliquot chain function, called <code>aliquot_chain2</code>, which takes only the argument <code>a</code> and which terminates not after a set number of iterations but when the sequence encounters a value it has encountered before. 

Test your function: <code>aliquot_chain2(24)</code> should return <code>[24, 36, 55, 17, 1, 0, 0]</code>, whereas <code>aliquot_chain2(28)</code> should return <code>[28, 28]</code>, and <code>aliquot_chain2(220)</code> should return <code>[220, 284, 220]</code>.

<!-- grade_id: cell-1b189024c6e59d88, points: 0, grade: false, solution: true -->
```python
def aliquot_chain2(a):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-425a664ab09929e4, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3b(_globals)
```

(c) Write and test a function called <code>fixed_point_list</code>, which takes as its input a function <code>f</code>, an initial value <code>x0</code> and a <code>tolerance</code>, and returns a list consisting of 

<code>[x0, f(x0), f(f(x0)), ...]</code>, 

which terminates when successive iterates are within <code>tolerance</code> of one another. (Make sure you test it only on iterations that you know converge; later, we'll look at how to build in safeguards to prevent infinite loops when an iteration fails to converge.)

<!-- grade_id: cell-70aad741029c978d, points: 0, grade: false, solution: true -->
```python
def fixed_point_list(f, x0, tolerance):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-753024b1ab5e656d, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3c(_globals)
```

## Question 4

The method of calculating the primes by successive filtering is called the <b>sieve of Eratosthenes</b>; Earlier, you explored an implementation of this based on <em>conditional comprehensions</em>. You started with the integers between 2 and 400 inclusive; you then removed all the multiples of 2 except 2 itself; then all the multiples of 3 except 3 itself; then all the multiples of 5 except 5 itself; and so on, using primes each time.

Write and test a function called <code>my_primes</code> that takes as its argument a single <code>n</code>, which you can assume to be an integer greater than 1, and, using the sieve of Eratosthenes, returns a list of all the primes between 2 and <code>n</code> inclusive. Make your function reasonably efficient; don't put in lots of unnecessary iterations.

<!-- grade_id: cell-09d7c2ad10004e0e, points: 0, grade: false, solution: true -->
```python
def my_primes(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c749e7603ceb65e9, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4(_globals)
```

## Bonus Questions

### Question 1

Write a function, `symCheck`, that takes as input a matrix and determines whether or not the matrix is symmetric. The function should return the strings `not symmetric` or `symmetric`. You may not use any obvious builtin function.

<!-- grade_id: cell-f3f2aefc41f2911d, points: 0, grade: false, solution: true -->
```python
# Question 1
def symCheck(A):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0ac2746c8242a091, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = [[1, 2, 3], [2, 1, 3]]
assert(symCheck(A) == 'not symmetric')
print('Test case passed!!!')
```

<!-- grade_id: cell-af2c44662743c74d, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = [[1, 2, 3], [2, 1, 3], [2, 1, 3]]
assert(symCheck(A) == 'not symmetric')
print('Test case passed!!!')
```

<!-- grade_id: cell-35974a0e424c744f, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = [[1, 2, 3], [2, 1, 3], [3, 3, 3]]
assert(symCheck(A) == 'symmetric')
print('Test case passed!!!')
```

<!-- grade_id: cell-3c74338d7cde78dc, points: 1, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 1 mark
```

<!-- grade_id: cell-6a3ea90ee9b77c8a, points: 1, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 1 mark
```

### Question 2

Write a function that takes in two arrays and determines all integers satisfying the following conditions:

* All elements of the first array are factors of the integer under consideration.
* The integer under consideration is a factor of all the elements in the second array.
    
Your function must return the amount of such numbers. For example, if $a = [2, 6]$ and $b = [24, 36]$, there are 2 numbers between them which satisfy the above, namely 6 and 12. To check, consider the first value 6: 6 % 2 = 0, 6 % 6 = 0, as well as 24 % 6 = 0 and 36 % 6 = 0. Similarly, 12 % 2 = 0, 12 % 6 = 0,  24 % 12 = 0 and 36 % 12 = 0.

<!-- grade_id: cell-c1ca2727aaac882f, points: 0, grade: false, solution: true -->
```python
def setDiff(a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-aca1bbc826cd86df, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 2 marks
# Unit test
bonus_q2(_globals)
```

<!-- grade_id: cell-bb1274933762d2a4, points: 1, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 1 mark
```

<!-- grade_id: cell-9f27c0f689941e99, points: 1, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 1 mark
```