```python
import time
from autograder import *
import math
from math import log, exp, cos, atan, sqrt, pi, e
import numpy as np
import cmath
_globals = globals()
# Worksheet answered by Mashrur Khondokar
```

Output:
```
Autograder loaded successfully!
Remember to always restart and run all from the Kernel menu before submitting!
```

# <center> Introduction to Mathematical Computing </center>
### <center>Phil Ramsden </center>

# <center> Worksheet 1: Arithmetic and functions </center>

**NOTE:** When answering a question, remember to remove the `raise NotImplementedError()` line.

## Question 1(i)

Calculate: 

(a) $4^2+1$;
(b) $4^{2+1}$;
(c) $2/3$ as a decimal;
(d) the largest integer smaller than $381/47$;
(e) $4 \times 2^5$;
(f) $(4 \times 2)^5$.

<!-- grade_id: cell-bcf9158f78bdb3a6, points: 0, grade: false, solution: true -->
```python
# Assign your solutions to the questions above to the below variables.
q1ia_answer = 4**2 + 1
q1ib_answer = 4**(2+1)
q1ic_answer = 2/3
q1id_answer = 381//47
q1ie_answer = 4*(2**5)
q1if_answer = (4*2)**5
print(q1ia_answer)
print(q1ib_answer)
print(q1ic_answer)
print(q1id_answer)
print(q1ie_answer)
print(q1if_answer)
# YOUR CODE HERE
```

Output:
```
17
64
0.6666666666666666
8
128
32768
```

<!-- grade_id: cell-b12f71f8ddaf1387, points: 6, grade: true, solution: false -->
```python
# 6 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1(_globals)
```

Output:
```
Question 1(i)a passed!
Question 1(i)b passed!
Question 1(i)c passed!
Question 1(i)d passed!
Question 1(i)e passed!
Question 1(i)f passed!
6 out of 6 tests passed
6
```

## Question 1(ii)

Calculate:

(a) The residue of $372517$ modulo $1820$.

(b) The residue of $(1417456 + 1191761)$ modulo $1820$.

(c) The residue of $(1417456 \times 1191761)$ modulo $1820$.

(d) The residue of $1417456^{1191761}$ modulo $1820$ (for best results use the <code>pow</code> function with three arguments).

<!-- grade_id: cell-593c9367c3c6115a, points: 0, grade: false, solution: true -->
```python
# Assign your solutions to the questions above to the below variables.
q1iia_answer = 372517 % 1820
q1iib_answer = (1417456 + 1191761) % 1820
q1iic_answer = (1417456*1191761) % 1820
start = time.time()
q1iid_answer = pow(1417456, 1191761, 1820)
duration = time.time() - start
print(q1iia_answer)
print(q1iib_answer)
print(q1iic_answer)
print(q1iid_answer)

# YOUR CODE HERE
```

Output:
```
1237
1157
636
976
```

<!-- grade_id: cell-d5ff3aef45b2ad85, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete `this cell
# Run this cell for grading
_globals = globals()
question1_ii(_globals)
```

Output:
```
Question 1(ii)a passed!
Question 1(ii)b passed!
Question 1(ii)c passed!
Question 1(ii)d passed!
5 out of 5 tests passed
5
```

## Question 2(i)

Using the <code>math</code> module, calculate: (a) $\sqrt{50}$; (b) $\cos \pi/7$; (c) $e^{e^2}$; (d) $\ln 58$; (e) $\log_{10} 58$; (f) $\cot^{-1} 2$; (g) $\coth^{-1} 2$.

<!-- grade_id: cell-c725db911b14d4ee, points: 0, grade: false, solution: true -->
```python
# begin by importing the necessary functions and constants from the math module
# Assign your solutions to the questions above to the below variables.
q2ia_answer = sqrt(50)
q2ib_answer = cos((pi/7))
q2ic_answer = exp(exp(2))
q2id_answer = log(58)
q2ie_answer = log(58, 10)
q2if_answer = atan(1/2)
q2ig_answer = 0.5*log((2+1)/(2-1))
print(q2ia_answer)
print(q2ib_answer)
print(q2ic_answer)
print(q2id_answer)
print(q2ie_answer)
print(q2if_answer)
print(q2ig_answer)

# YOUR CODE HERE
```

Output:
```
7.0710678118654755
0.9009688679024191
1618.1779919126539
4.060443010546419
1.7634279935629371
0.4636476090008061
0.5493061443340549
```

<!-- grade_id: cell-660c3f636dc0be36, points: 7, grade: true, solution: false -->
```python
# 7 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2_i(_globals)
```

Output:
```
Question 2(i)a passed!
Question 2(i)b passed!
Question 2(i)c passed!
Question 2(i)d passed!
Question 2(i)e passed!
Question 2(i)f passed!
Question 2(i)g passed!
7 out of 7 tests passed
7
```

## Question 2(ii)

The <code>numpy</code> module contains many of the same functions as <code>math</code> (though they're slightly different, in ways we'll explore). Recalculate (a) - (g) using this module instead.

<!-- grade_id: cell-45ed8f2561c9f720, points: 0, grade: false, solution: true -->
```python
# begin by importing the necessary functions and constants from the math module
# Assign your solutions to the questions above to the below variables.
q2iia_answer = np.sqrt(50)
q2iib_answer = np.cos(np.pi/7)
q2iic_answer = np.exp(np.exp(2))
q2iid_answer = np.log(58)
q2iie_answer = np.log10(58)
q2iif_answer = np.arctan(1/2)
q2iig_answer = 0.5*np.log((2+1)/(2-1))
print(q2iia_answer)
print(q2iib_answer)
print(q2iic_answer)
print(q2iid_answer)
print(q2iie_answer)
print(q2iif_answer)
print(q2iig_answer)

# YOUR CODE HERE
```

Output:
```
7.0710678118654755
0.9009688679024191
1618.1779919126539
4.060443010546419
1.7634279935629373
0.4636476090008061
0.5493061443340549
```

<!-- grade_id: cell-6390293366653fbb, points: 7, grade: true, solution: false -->
```python
# 7 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2_ii(_globals)
```

Output:
```
Question 2(ii)a passed!
The cosine function does not seem to come from the numpy module
Question 2(ii)b failed!
Question 2(ii)c passed!
The logarithm function does not seem to come from the numpy module
Question 2(ii)d failed!
Question 2(ii)e passed!
Question 2(ii)f passed!
Question 2(ii)g passed!
5 out of 7 tests passed
5
```

## Question 3

(a) Define variables $p=6$ and $q=9$, and calculate the quantity
$$x=\left(\frac{q}{2}+\sqrt{\left[\frac{q}{2}\right]^2-\left[\frac{p}{3}\right]^3}\right)^{1/3}+\left(\frac{q}{2}-\sqrt{\left[\frac{q}{2}\right]^2-\left[\frac{p}{3}\right]^3}\right)^{1/3}.$$
Verify Cardano's Theorem, that $x$ is a solution of the cubic
$$x^3=p\,x+q.$$

<!-- grade_id: cell-05998432c3deefe2, points: 0, grade: false, solution: true -->
```python
# import any necessary functions
# Assign your solutions to the questions above to the below variables.
p = 6
q = 9
q2 = q/2
p3 = p/3

x = (q2 + sqrt(q2**2 - p3**3))**(1/3) + (q2 - sqrt(q2**2 - p3**3))**(1/3)
print(p)
print(q)
print(x)
# YOUR CODE HERE
```

Output:
```
6
9
3.0
```

<!-- grade_id: cell-721979068553d977, points: 1, grade: true, solution: false -->
```python
#1 Mark
assert(x == question3(p, q))
print("Question 3(a) passed!")
```

Output:
```
Question 3(a) passed!
```

(b) Repeat for $p=3$ and $q=18$.

<!-- grade_id: cell-972547089f6b9cc6, points: 0, grade: false, solution: true -->
```python
# import any necessary functions
# Assign your solutions to the questions above to the below variables.
p = 3 
q = 18
q2 = q/2
p3 = p/3

x = (q2 + sqrt(q2**2 - p3**3))**(1/3) + (q2 - sqrt(q2**2 - p3**3))**(1/3)
print(p)
print(q)
print(x)
# YOUR CODE HERE
```

Output:
```
3
18
2.999999999999999
```

<!-- grade_id: cell-22cd955862519f1c, points: 1, grade: true, solution: false -->
```python
#1 Mark
assert(x == question3(p, q))
print("Question 3(b) passed!")
```

Output:
```
Question 3(b) passed!
```

(c) If you try to repeat for $p=6$ and $q=4$, you may run into a problem. Can you solve it?

<!-- grade_id: cell-d7a2fd2365c71adf, points: 0, grade: false, solution: true -->
```python
# import any necessary functions
# Assign your solutions to the questions above to the below variables.
p = 6
q = 4
q2 = q/2
p3 = p/3

x = (q2 + cmath.sqrt(q2**2 - p3**3))**(1/3) + (q2 - cmath.sqrt(q2**2 - p3**3))**(1/3)
print(p)
print(q)
print(x)
# YOUR CODE HERE
```

Output:
```
6
4
(2.7320508075688776+0j)
```

<!-- grade_id: cell-69e5c37ef72d2e4e, points: 1, grade: true, solution: false -->
```python
#1 Mark
assert(x == question3(p, q))
print("Question 3(c) passed!")
```

Output:
```
Question 3(c) passed!
```

## Question 4

Experiment with <b>integer division</b>, and answer the following question:

For integers $a$ and $b$, the code

<code>a // b</code>

always returns

1. The least integer greater than or equal to $a/b$;
2. The greatest integer less than or equal to $a/b$;
3. The nearest integer to $a/b$;
4. The number $a/b$, rounded towards zero;
5. The number $a/b$, rounded away from zero;
6. None of these; it depends on the signs of $a$ and $b$.

<!-- grade_id: cell-87d8b531ab5cfff5, points: 0, grade: false, solution: true -->
```python
#q4_answer = # enter one of 1, 2, 3, 4, 5 or 6
# Assign your solutions to the questions above to the below variables.
# YOUR CODE HERE
q4_answer = 2
```

<!-- grade_id: cell-2c3657af708f4b51, points: 1, grade: true, solution: false -->
```python
#1 Mark
assert(q4_answer == question4())
print("Question 4 passed!")
```

Output:
```
Question 4 passed!
```

## Question 5 (i)

The expressions (a) $\sin^{-1}(-2)$; (b) $\cos^{-1}(-2)$; (c) $\tanh^{-1}(-2)$; (d) $\ln(-2)$

can all be given sensible mathematical values. Using functions from a suitable Python module, calculate these values.

<!-- grade_id: cell-8de18ecb33e6d43d, points: 0, grade: false, solution: true -->
```python
# Assign your solutions to the questions above to the below variables.

q5ia_answer = cmath.asin(-2)
q5ib_answer = cmath.acos(-2)
q5ic_answer = cmath.atanh(-2)
q5id_answer = cmath.log(-2)

print(q5ia_answer)
print(q5ib_answer)
print(q5ic_answer)

# YOUR CODE HERE
```

Output:
```
(-1.5707963267948966+1.3169578969248166j)
(3.141592653589793-1.3169578969248166j)
(-0.5493061443340549+1.5707963267948966j)
```

<!-- grade_id: cell-68368eae35986782, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question5_i(_globals)
```

Output:
```
Question 5(i)a passed!
Question 5(i)b passed!
Question 5(i)c passed!
Question 5(i)d passed!
4 out of 4 tests passed
4
```

## Question 5 (ii)

The expressions (a) $\sin^{-1}(-2)$; (b) $\cos^{-1}(-2)$; (c) $\tanh^{-1}(-2)$; (d) $\ln(-2)$

can actually all be calculated using only functions from the <code>numpy</code> module, though there's a bit of a trick to making it work. Find the trick, and make it work!

<!-- grade_id: cell-966f608a75b6df3f, points: 0, grade: false, solution: true -->
```python
# Assign your solutions to the questions above to the below variables.
q5iia_answer = np.arcsin(np.complex(-2))
q5iib_answer = np.arccos(np.complex(-2))
q5iic_answer = np.arctanh(np.complex(-2))
q5iid_answer = np.log(np.complex(-2))
print(q5iia_answer)
print(q5iib_answer)
print(q5iic_answer)
print(q5iid_answer)
```

Output:
```
(-1.5707963267948966+1.3169578969248166j)
(3.141592653589793-1.3169578969248166j)
(-0.5493061443340549+1.5707963267948966j)
(0.6931471805599453+3.141592653589793j)
```

<!-- grade_id: cell-c794ccf9d2ca10bb, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
question5_ii(_globals)
```

Output:
```
Question 5(ii)a passed!
Question 5(ii)b passed!
Question 5(ii)c passed!
Question 5(ii)d passed!
4 out of 4 tests passed
4
```

```python

```