```python
from autograder import *
import math
_globals = globals()
```

Output:
```
Autograder loaded successfully!
Remember to always restart and run all from the Kernel menu before submitting!
```

# <center> Introduction to Mathematical Computing </center>
### <center> Phil Ramsden </center>

# <center> Worksheet 2: Data types and data structures </center>

**NOTE:** When answering a question, remember to remove the `raise NotImplementedError()` line.

## Question 1(i)

(a) Define ints <code>int1</code>, <code>int2</code> and <code>int3</code> equal to, respectively, 123, 456 and 789.

<!-- grade_id: cell-f6b69f6c0f33441a, points: 0, grade: false, solution: true -->
```python
int1 = 123
int2 = 456
int3 = 789
```

<!-- grade_id: cell-54797219db211091, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1_i_a(_globals)# 3 Marks

```

(b) Define <b>floats</b> <code>float1</code>, <code>float2</code> and <code>float3</code> with the same numerical values as <code>int1</code>, <code>int2</code> and <code>int3</code>, respectively.

<!-- grade_id: cell-b54f78cd269916a2, points: 0, grade: false, solution: true -->
```python
float1 = float(123)
float2 = float(456)
float3 = float(789)
```

<!-- grade_id: cell-0ebff7f8b80df518, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1_i_b(_globals)
```

(c) By experimenting with your six variables, answer the following multiple-choice question.

If <code>x</code> is an int, and <code>y</code> is a float with the same numerical value, then 

<ol>
    <li>the command <code>x == y</code> returns <code>True</code>, and <code>type(x)</code> and <code>type(y)</code> have the same value;</li>
    <li>the command <code>x == y</code> returns <code>False</code>, and <code>type(x)</code> and <code>type(y)</code> have different values;</li>
    <li>the command <code>x == y</code> returns <code>True</code>, but <code>type(x)</code> and <code>type(y)</code> have different values;</li>
    <li>the command <code>x == y</code> returns <code>False</code>, but <code>type(x)</code> and <code>type(y)</code> have the same value.</li>
</ol>

Fill in the correct numerical value, 1, 2, 3 or 4, below.

<!-- grade_id: cell-e5759f3ce7494418, points: 0, grade: false, solution: true -->
```python
question1ic_answer = 3
print(int1==float1)
```

<!-- grade_id: cell-4ac01a582155b515, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question1ic_answer == question1_i_c())
print('test case passed!')
```

(d) By experimenting with adding, subtracting, multiplying and dividing your variables, answer the following multiple-choice question.

If we add, subtract, multiply or divide an int and a float, the result is

<ol>
    <li>always an int;</li>
    <li>always a float;</li>
    <li>an int or a float, depending on the operation we choose;</li>
    <li>an int or a float, depending on whether the int or the float comes first in the expression.</li>
</ol>

Fill in the correct numerical value, 1, 2, 3 or 4, below.

<!-- grade_id: cell-5479e40c5e2b9e62, points: 0, grade: false, solution: true -->
```python
question1id_answer = 2
experiment = float1/int1
print(type(experiment))
```

<!-- grade_id: cell-33e61bed96005469, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question1id_answer == question1_i_d())
print('test case passed!')
```

(e) By experimenting with raising one of your variables to the power of another, choose the true statement from the following:

<ol>
    <li>Large ints are represented as exact quantities in Python, no matter how many digits they have.</li>
    <li>Ints above a certain size have to be represented in the form $a\times10^r$, but if they're very large, there's an overflow error.</li>
    <li>Ints above a certain size have to be represented in the form $a\times10^r$, with no danger of an overflow error.</li>
</ol>

Fill in the correct numerical value, 1, 2 or 3, below.

<!-- grade_id: cell-c07c6f59097c20c6, points: 0, grade: false, solution: true -->
```python
question1ie_answer = 1
bigint = 1500**456
print(bigint)
```

<!-- grade_id: cell-8bc067badf04cc2c, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question1ie_answer == question1_i_e())
print('test case passed!')
```

(f) By experimenting with raising one of your variables to the power of another, choose the true statement from the following:

<ol>
    <li>Large floats are represented as exact quantities in Python, no matter how many digits they have.</li>
    <li>Floats above a certain size have to be represented in the form $a\times10^r$, but if they're very large, there's an overflow error.</li>
    <li>Floats above a certain size have to be represented in the form $a\times10^r$, with no danger of an overflow error.</li>
</ol>

Fill in the correct numerical value, 1, 2 or 3, below.

<!-- grade_id: cell-de4b5e500d5a3858, points: 0, grade: false, solution: true -->
```python
question1if_answer = 2
```

<!-- grade_id: cell-fb8aa4e27ae298fb, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question1if_answer == question1_i_f())
print('test case passed!')
```

(g) Import the <code>math</code> module, and experiment with mathematical functions of ints and floats. Answer the following question.

The sine, cosine, exponential and logarithmic functions in the <code>math</code> module:

<ol>
    <li>all return an int if the argument is an int, and a float if the argument is a float;</li>
    <li>all always return an int;</li>
    <li>all always return a float;</li>
    <li>behave differently from one another.</li>
</ol>

Fill in the correct numerical value, 1, 2, 3 or 4, below.

<!-- grade_id: cell-a63c3f2e5f3a1b41, points: 0, grade: false, solution: true -->
```python
question1ig_answer = 3
print(type(math.cos(30)), type(math.sin(30)), type(math.log(30)), type(math.exp(30)))
```

<!-- grade_id: cell-f6f2200d21835569, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question1ig_answer == question1_i_g())
print('test case passed!')
```

## Question 1(ii)

(a) Define the complex numbers <code>complex1</code>, <code>complex2</code> and <code>complex3</code> with the same numerical values as, respectively, <code>int1</code>, <code>int2</code> and <code>int3</code>.

Using the <code>complex</code> constructor function, define <code>complex4</code>, using <code>int1</code> and <code>int2</code> as real and imaginary part, and <code>complex5</code>, using <code>float1</code> and <code>float2</code> as real and imaginary part,

<!-- grade_id: cell-ea86d15e54090326, points: 0, grade: false, solution: true -->
```python
complex1 = int1 + 0j
complex2 = int2 + 0j
complex3 = int3 + 0j
complex4 = complex(int1, int2)
complex5 = complex(float1, float2)

```

<!-- grade_id: cell-d7d5868f5f17ec70, points: 10, grade: true, solution: false -->
```python
# 10 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1_ii_a(_globals)
```

(b) By experimenting, answer the following multiple-choice question.

If <code>x</code> is an int, and <code>z</code> is a complex with the same numerical value, then 

<ol>
    <li>the command <code>x == z</code> returns <code>True</code>, and <code>type(x)</code> and <code>type(z)</code> have the same value;</li>
    <li>the command <code>x == z</code> returns <code>False</code>, and <code>type(x)</code> and <code>type(z)</code> have different values;</li>
    <li>the command <code>x == z</code> returns <code>True</code>, but <code>type(x)</code> and <code>type(z)</code> have different values;</li>
    <li>the command <code>x == z</code> returns <code>False</code>, but <code>type(x)</code> and <code>type(z)</code> have the same value.</li>
</ol>

Fill in the correct numerical value, 1, 2, 3 or 4, below.

<!-- grade_id: cell-0b6aca3c14ba42ba, points: 0, grade: false, solution: true -->
```python
question1iib_answer = 3
```

<!-- grade_id: cell-cee4552ab0535602, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question1iib_answer == question1_ii_b())
print('test case passed!')
```

(c) By experimenting, answer the following multiple-choice question.

If <code>y</code> is a float, and <code>z</code> is a complex with the same numerical value, then 

<ol>
    <li>the command <code>y == z</code> returns <code>True</code>, and <code>type(y)</code> and <code>type(z)</code> have the same value;</li>
    <li>the command <code>y == z</code> returns <code>False</code>, and <code>type(y)</code> and <code>type(z)</code> have different values;</li>
    <li>the command <code>y == z</code> returns <code>True</code>, but <code>type(y)</code> and <code>type(z)</code> have different values;</li>
    <li>the command <code>y == z</code> returns <code>False</code>, but <code>type(y)</code> and <code>type(z)</code> have the same value.</li>
</ol>

Fill in the correct numerical value, 1, 2, 3 or 4, below.

<!-- grade_id: cell-851a99484184f3d8, points: 0, grade: false, solution: true -->
```python
question1iic_answer = 3
```

<!-- grade_id: cell-cb59d57abf447995, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question1iic_answer == question1_ii_c())
print('test case passed!')
```

(d) By experimenting, answer the following multiple-choice question.

If we add, subtract, multiply or divide an int and a complex, the result is

<ol>
    <li>always an int;</li>
    <li>always a complex;</li>
    <li>an int or a complex, depending on the operation we choose;</li>
    <li>an int or a complex, depending on whether the int or the complex comes first in the expression.</li>
</ol>

Fill in the correct numerical value, 1, 2, 3 or 4, below.



<!-- grade_id: cell-21c341b2ccad6e5c, points: 0, grade: false, solution: true -->
```python
question1iid_answer = 2
```

<!-- grade_id: cell-4aa901bc8a543643, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question1iid_answer == question1_ii_d())
print('test case passed!')
```

(e) By experimenting, answer the following multiple-choice question.

If we add, subtract, multiply or divide a float and a complex, the result is

<ol>
    <li>always a float;</li>
    <li>always a complex;</li>
    <li>a float or a complex, depending on the operation we choose;</li>
    <li>a float or a complex, depending on whether the float or the complex comes first in the expression.</li>
</ol>

Fill in the correct numerical value, 1, 2, 3 or 4, below.

<!-- grade_id: cell-020ae20e70e4982a, points: 0, grade: false, solution: true -->
```python
question1iie_answer = 2
print(type(5.3/3+1j), type((5.0+1j)-3.0), type(3.2+(4.5+0.5j)))
```

<!-- grade_id: cell-284978c168e50d93, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question1iie_answer == question1_ii_e())
print('test case passed!')
```

(f) By experimenting with commands like

```python
complex1.real
```

and

```python
complex1.imag
```

answer the following multiple choice question.

The real and imaginary parts of a complex are

<ol>
    <li>always ints;</li>
    <li>always floats;</li>    
    <li>sometimes ints and sometimes floats.</li>
</ol>

Fill in the correct numerical value, 1, 2 or 3, below.

<!-- grade_id: cell-647289a9bef34648, points: 0, grade: false, solution: true -->
```python
question1iif_answer = 2
comp1 = 5+3j
print(type(comp1.real), type(comp1.real))
```

<!-- grade_id: cell-06ca8da159c442f9, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question1iif_answer == question1_ii_f())
print('test case passed!')
```

(g) Experiment with doing arithmetic with `True` and `False`, via code like
```python
2 + True
5 * False
7 ** False
exp(True)
```

Which of the following statements is true?

<ol>
    <li>It is not possible to do arithmetic with Boolean data;</li>
    <li>It is possible to do arithmetic with Boolean data, and <code>False</code> and <code>True</code> are treated like <code>0</code> and <code>1</code> respectively;</li>
    <li>It is possible to do arithmetic with Boolean data, and <code>False</code> and <code>True</code> are treated like <code>-1</code> and <code>1</code> respectively.</li>
</ol>

Fill in the correct numerical value, 1, 2 or 3, below.

<!-- grade_id: cell-83aa3cd1875fed96, points: 0, grade: false, solution: true -->
```python
question1iig_answer = 2
print(2+True, 5*False, 7**False, math.exp(True))
```

<!-- grade_id: cell-26cf1638c93c2fb9, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question1iig_answer == question1_ii_g())
print('test case passed!')
```

(h) Investigate the `^` operator when used on Boolean data, via code such as
```python
True ^ False
3<4 ^ 2==2
``` 
and so on. Which of the following is true?

<ol>
    <li>For any Boolean expressions <code>p</code> and <code>q</code>, <code>p ^ q</code> is the same as <code>p or q</code>;</li>
    <li>For any Boolean expressions <code>p</code> and <code>q</code>, <code>p ^ q</code> is the same as <code>p and q</code>;</li>
    <li>For any Boolean expressions <code>p</code> and <code>q</code>, <code>p ^ q</code> is the same as <code>not (p or q)</code>;</li>
    <li>For any Boolean expressions <code>p</code> and <code>q</code>, <code>p ^ q</code> is the same as <code>not (p and q)</code>;</li>
    <li>For any Boolean expressions <code>p</code> and <code>q</code>, <code>p ^ q</code> is the same as <code>(p or q) and not (p and q)</code>.</li>
</ol>

Fill in the correct numerical value, 1, 2, 3, 4 or 5, below.

<!-- grade_id: cell-2ebab91f303381ef, points: 0, grade: false, solution: true -->
```python
question1iih_answer = 5
print(True^False)
```

<!-- grade_id: cell-7a82b8797fa34f56, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question1iih_answer == question1_ii_h())
print('test case passed!')
```

(i) Investigate the `^` operator when used with ints, via code such as
```python
4 ^ 5
4 ^ 8
```
and so on. Write down a short description of what it does. (Autograding can't help here.)

```python
# The ^ operator is a binary xor (exclusive or) operator. So when used in ints (say a and b),
# it returns the number corresponding to the binary translation of the binary number that
# has the bits in a or b, but not ones in both.
```

## Question 2(i)

For the given string <code>string1</code>, calculate

(a) A string containing all the characters of <code>string1</code> between the first and the 15th inclusive.

(b) A string containing the first, third, fifth characters, etc, of <code>string1</code>, up to the 15th, inclusive.

(c) A version of <code>string1</code> entirely in upper case.

(d) A single string consisting of three copies of <code>string1</code> one after the other. 

<!-- grade_id: cell-2d7d7363cf1d2a85, points: 0, grade: false, solution: true -->
```python
string1 = 'The quick brown fox jumps over the lazy dog'

q2ia_answer = string1[0:15]
q2ib_answer = string1[0:16:2]
q2ic_answer = string1.upper()
q2id_answer = 3*string1
print(q2ia_answer)
print(q2ib_answer)
print(q2ic_answer)
print(q2id_answer)
```

<!-- grade_id: cell-0708126b0ea18264, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2_i(_globals)
```

## Question 2(ii)

For the given list <code>list1</code>, calculate

(a) A list containing all the elements of <code>list1</code> between the first and the 11th inclusive.

(b) A list containing the first, third, fifth elements, etc, of <code>list1</code>, up to the 11th, inclusive.

(c) A single list consisting of three copies of <code>list1</code>, one after the other. 

<!-- grade_id: cell-4b61ddd6fa75232e, points: 0, grade: false, solution: true -->
```python
list1 = [0+0j, 'one', 2, 3.0, 4+0j, 'five', 6, 7.0, 8+0j, 'nine', 10, 11.0, 12+0j, 'thirteen', 14, 15.0]
q2iia_answer = list1[0:11]
q2iib_answer = list1[0:12:2]
q2iic_answer = 3*list1
print(q2iia_answer)
print(q2iib_answer)
print(q2iic_answer)
```

Output:
```
[0j, 'one', 2, 3.0, (4+0j), 'five', 6, 7.0, (8+0j), 'nine', 10]
[0j, 2, (4+0j), 6, (8+0j), 10]
[0j, 'one', 2, 3.0, (4+0j), 'five', 6, 7.0, (8+0j), 'nine', 10, 11.0, (12+0j), 'thirteen', 14, 15.0, 0j, 'one', 2, 3.0, (4+0j), 'five', 6, 7.0, (8+0j), 'nine', 10, 11.0, (12+0j), 'thirteen', 14, 15.0, 0j, 'one', 2, 3.0, (4+0j), 'five', 6, 7.0, (8+0j), 'nine', 10, 11.0, (12+0j), 'thirteen', 14, 15.0]
```

<!-- grade_id: cell-1b662e29550d694b, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2_ii(_globals)
```

Output:
```
All tests passed!!
3
```

## Question 2(iii)

For the given tuple <code>tuple1</code>, calculate

(a) A tuple containing all the elements of <code>tuple1</code> between the first and the 11th inclusive.

(b) A tuple containing the first, third, fifth elements, etc, of <code>tuple1</code>, up to the 11th, inclusive.

(c) A single tuple consisting of three copies of <code>tuple1</code>, one after the other. 

<!-- grade_id: cell-b8aa41567b78e5e8, points: 0, grade: false, solution: true -->
```python
tuple1 = (0+0j, 'one', 2, 3.0, 4+0j, 'five', 6, 7.0, 8+0j, 'nine', 10, 11.0, 12+0j, 'thirteen', 14, 15.0)

q2iiia_answer = tuple1[0:11]
q2iiib_answer = tuple1[0:12:2]
q2iiic_answer = 3*tuple1
print(q2iiia_answer)
print(q2iiib_answer)
print(q2iiic_answer)
# YOUR CODE HERE
```

Output:
```
(0j, 'one', 2, 3.0, (4+0j), 'five', 6, 7.0, (8+0j), 'nine', 10)
(0j, 2, (4+0j), 6, (8+0j), 10)
(0j, 'one', 2, 3.0, (4+0j), 'five', 6, 7.0, (8+0j), 'nine', 10, 11.0, (12+0j), 'thirteen', 14, 15.0, 0j, 'one', 2, 3.0, (4+0j), 'five', 6, 7.0, (8+0j), 'nine', 10, 11.0, (12+0j), 'thirteen', 14, 15.0, 0j, 'one', 2, 3.0, (4+0j), 'five', 6, 7.0, (8+0j), 'nine', 10, 11.0, (12+0j), 'thirteen', 14, 15.0)
```

<!-- grade_id: cell-e604802244cecd7a, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2_iii(_globals)
```

Output:
```
All tests passed!!
3
```

## Question 2(iv)

(a) By means of a single use of the `extend` method (which you should look up), append the elements `16+0j`, `'seventeen'`, `18` and `19.0` to `list1`, so that the value of `list1` is now
```python
[0+0j, 'one', 2, 3.0, 4+0j, 'five', 6, 7.0, 8+0j, 'nine', 10, 11.0, 12+0j, 'thirteen', 14, 15.0, 16+0j, 'seventeen', 18, 19.0]
```
Autograding won't be able to tell whether you've used `extend`, and there are other ways of doing it, but please try to do so!

<!-- grade_id: cell-8e4ccb1216d7534d, points: 0, grade: false, solution: true -->
```python
list1.extend([16+0j, 'seventeen', 18, 19.0])

print(list1)

```

Output:
```
[0j, 'one', 2, 3.0, (4+0j), 'five', 6, 7.0, (8+0j), 'nine', 10, 11.0, (12+0j), 'thirteen', 14, 15.0, (16+0j), 'seventeen', 18, 19.0]
```

<!-- grade_id: cell-d2633add38b13ebb, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(list1 == question2_iv_a())
print('test case passed!')
```

Output:
```
test case passed!
```

(b) Change the five strings in `list1` to `'uno'`, `'cinco'`, `'nueve'`, `'trece'` and `'diecisiete'` respectively. Try to do this with a <b>single command</b>.

<!-- grade_id: cell-6fd9bacae9d3f248, points: 0, grade: false, solution: true -->
```python
list1[1:18:4] = 'uno', 'cinco', 'nueve', 'trece', 'diecisiete'

print(list1)
```

Output:
```
[0j, 'uno', 2, 3.0, (4+0j), 'cinco', 6, 7.0, (8+0j), 'nueve', 10, 11.0, (12+0j), 'trece', 14, 15.0, (16+0j), 'diecisiete', 18, 19.0]
```

<!-- grade_id: cell-16f100ca6ae8e956, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(list1 == question2_iv_b())
print('test case passed!')
```

Output:
```
test case passed!
```

(c) Try to find a way to do the equivalent of part (a) with `tuple1`, so that `tuple1` becomes
```python
(0+0j, 'one', 2, 3.0, 4+0j, 'five', 6, 7.0, 8+0j, 'nine', 10, 11.0, 12+0j, 'thirteen', 14, 15.0, 16+0j, 'seventeen', 18, 19.0)
```
There's no terribly elegant way to do this, but if you can, avoid simply typing out (or, worse, copying and pasting) all twenty elements!

<!-- grade_id: cell-f35cddbe5f0f0eee, points: 0, grade: false, solution: true -->
```python
list3 = list(tuple1)
list3.extend([16+0j, 'seventeen', 18, 19.0])
print(list3)
tuple1 = tuple(list3)
print(tuple1)
```

Output:
```
[0j, 'one', 2, 3.0, (4+0j), 'five', 6, 7.0, (8+0j), 'nine', 10, 11.0, (12+0j), 'thirteen', 14, 15.0, (16+0j), 'seventeen', 18, 19.0]
(0j, 'one', 2, 3.0, (4+0j), 'five', 6, 7.0, (8+0j), 'nine', 10, 11.0, (12+0j), 'thirteen', 14, 15.0, (16+0j), 'seventeen', 18, 19.0)
```

<!-- grade_id: cell-ecd565c6e8540ff4, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(tuple1 == question2_iv_c())
print('test case passed!')
```

Output:
```
test case passed!
```

(d) Try to find a way to do the equivalent of part (b) with `tuple1`, so that `tuple1` becomes
```python
(0j, 'uno', 2, 3.0, (4+0j), 'cinco', 6, 7.0, (8+0j), 'nueve', 10, 11.0, (12+0j), 'trece', 14, 15.0, (16+0j), 'diecisiete', 18, 19.0)
```
Again, try to avoid simply typing out (or, worse, copying and pasting) all twenty elements.

<!-- grade_id: cell-3b1228a905d86fbc, points: 0, grade: false, solution: true -->
```python
print(tuple1)
list4 = list(tuple1)
list4[1:18:4] = 'uno', 'cinco', 'nueve', 'trece', 'diecisiete'
tuple1 = tuple(list4)
```

Output:
```
(0j, 'one', 2, 3.0, (4+0j), 'five', 6, 7.0, (8+0j), 'nine', 10, 11.0, (12+0j), 'thirteen', 14, 15.0, (16+0j), 'seventeen', 18, 19.0)
```

<!-- grade_id: cell-bb186a18e242fcb4, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(tuple1 == question2_iv_d())
print('test case passed!')
```

Output:
```
test case passed!
```

## Question 3

(a) For the given string <code>string1</code>, use the <code>split</code> method to create a list <code>words1</code> consisting of all the words of <code>string1</code> in order.

<!-- grade_id: cell-77c4f8c0334919ba, points: 0, grade: false, solution: true -->
```python
string1 = "The quick brown fox jumps over the lazy dog"

words1 = string1.split(" ")
```

<!-- grade_id: cell-816895ca5f596490, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(words1 == question3_a())
print('test case passed!')
```

Output:
```
test case passed!
```

(b) Use the <code>join</code> method to create, from <code>words1</code>, a string <code>string2</code> identical to <code>string1</code>.

<!-- grade_id: cell-2bd5ec6af751e194, points: 0, grade: false, solution: true -->
```python
string2 = " ".join(words1)
```

<!-- grade_id: cell-ff5c25cf76611791, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(string2 == question3_b())
print('test case passed!')
```

Output:
```
test case passed!
```

(c) Use the <code>lower</code> method, the <code>split</code> method, the <code>join</code> method and the <code>sort</code> method to create a string <code>string3</code> consisting of all the words in <code>string1</code>, rendered in lower case, in alphabetical order. Words should be separated by a comma and a space, like so:

```python
'brown, dog, ...'
```

<!-- grade_id: cell-8be204034cff30b1, points: 0, grade: false, solution: true -->
```python
words2 = string1.lower().split(" ")
words2.sort()
string3 = ", ".join(words2)
```

<!-- grade_id: cell-1328bc4b570e5f94, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(string3 == question3_c())
print('test case passed!')
```

Output:
```
test case passed!
```

(d) Use the <code>replace</code> method to create a string <code>string4</code>, which is identical to <code>string1</code>, except that the word 'dog' has been replaced with the word 'cat'.

<!-- grade_id: cell-0474845fee817e64, points: 0, grade: false, solution: true -->
```python
string4 = string1.replace("dog", "cat")
```

<!-- grade_id: cell-8b747e7833b7a9e3, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(string4 == question3_d())
print('test case passed!')
```

Output:
```
test case passed!
```

(e) Use the <code>replace</code> method to create a string <code>string5</code>, which is identical to <code>string1</code>, except that the words 'fox' and 'dog' have been swapped. (This may need a little thought.)

<!-- grade_id: cell-1ac7e30bd9651e4e, points: 0, grade: false, solution: true -->
```python
string5 = string1.replace("dog", "fox"); string5 = string5.replace("fox", "dog", 1)
print(string5)
```

Output:
```
The quick brown dog jumps over the lazy fox
```

<!-- grade_id: cell-d26a3fdceb860e6c, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(string5 == question3_e())
print('test case passed!')
```

Output:
```
test case passed!
```

## Question 4

The following code uses a template, and the <code>format</code> method, to create a string, <code>polyhedron1</code>, that represents a true statement about the tetrahedron.

```python
polyhedron_template = 'The regular solid known as the {} has {} {} faces, {} edges and {} vertices.'

polyhedron1 = polyhedron_template.format('tetrahedron', 4, 'triangular', 6, 4)

print(polyhedron1)
```

Output:
```
The regular solid known as the tetrahedron has 4 triangular faces, 6 edges and 4 vertices.
```

(a) Using the same template, but different arguments to the <code>format</code> method, create strings <code>polyhedron2</code> and <code>polyhedron3</code> representing true statements about the cube and the octahedron respectively.

<!-- grade_id: cell-900127b037e04dbb, points: 0, grade: false, solution: true -->
```python
pht = polyhedron_template
polyhedron2 = pht.format("cube", 6, 'square', 12, 8)
polyhedron3 = pht.format("octahedron", 8, 'triangular', 12, 6)
print(polyhedron2)
print(polyhedron3)
```

Output:
```
The regular solid known as the cube has 6 square faces, 12 edges and 8 vertices.
The regular solid known as the octahedron has 8 triangular faces, 12 edges and 6 vertices.
```

<!-- grade_id: cell-0373209931e68134, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4_a(_globals)
```

Output:
```
2 tests passed!
2
```

(b) The data we need to use in a call to <code>format</code> doesn't always come in a convenient order. Write a new template called <code>polyhedron_template2</code>, written so that the input
```python
polyhedron_template2.format(8,12,6,'square','cube')
```
creates a string identical to <code>polyhedron2</code>. Test it on the other two examples.

<!-- grade_id: cell-98df2ff7aaa91555, points: 0, grade: false, solution: true -->
```python
polyhedron_template2 = 'The regular solid known as the {4} has {2} {3} faces, {1} edges and {0} vertices.'
polyhedron_template2.format(8,12,6,'square','cube')
```

Output:
```
'The regular solid known as the cube has 6 square faces, 12 edges and 8 vertices.'
```

<!-- grade_id: cell-c330f1df1757c6bf, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(polyhedron_template2.format(8,12,6,'square','cube') == question4_b())
print('test case passed!')
```

Output:
```
test case passed!
```

(c) Sometimes we need to work with <em>lists</em> of data. Write a third template called <code>polyhedron_template3</code>, written so that the following code...
```python
cube_data = [8, 12, 6, 'square', 'cube']
polyhedron_template3.format(cube_data)
```
... creates a string identical to <code>polyhedron2</code>. Test it on the other two examples.

<!-- grade_id: cell-404d63938f8efa3f, points: 0, grade: false, solution: true -->
```python
polyhedron_template3 = 'The regular solid known as the {0[4]} has {0[2]} {0[3]} faces, {0[1]} edges and {0[0]} vertices.'
cube_data = [8, 12, 6, 'square', 'cube']
polyhedron_template3.format(cube_data)
```

Output:
```
'The regular solid known as the cube has 6 square faces, 12 edges and 8 vertices.'
```

<!-- grade_id: cell-9ea07bfd55d9cb07, points: 1, grade: true, solution: false -->
```python
# 1 Mark
cube_data = [8, 12, 6, 'square', 'cube']
assert(polyhedron_template3.format(cube_data) == question4_c())
print('test case passed!')
```

Output:
```
test case passed!
```

```python

```